<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/The-Real-and-Complex-Number-Systems-(NS)/Extraction-of-roots/Complex-Root-of-a-Cubic-Equation.xml">
  <question type="stack">
    <name>
      <text>Complex Root of a Cubic Equation</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Determine in the form \(p + {@q*%i@}\) a complex root of the cubic equation \[ {@simpcubic@} = 0.\] &lt;/p&gt;
&lt;p&gt;Note: &lt;/p&gt;
&lt;p&gt;&lt;em&gt; - Recall that \(\sqrt{a} \) is  also rewritten as \( a^{@1/2@}\),  with \(a&gt;0\)  &lt;/em&gt;&lt;/p&gt;
&lt;p&gt;&lt;em&gt; - Type your answer in one expression using STACK notation, For example: 2^(3/2) + i*2^(3/2)&lt;/em&gt;&lt;/p&gt;
&lt;p&gt;&lt;em&gt; - Do not separate expressions with commas or “or.”&lt;/em&gt;&lt;/p&gt;
&lt;p&gt;&lt;/p&gt;
&lt;br /&gt; [[input:ans1]] [[validation:ans1]][[feedback:Result]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;&lt;b&gt;Worked Solution&lt;/b&gt;&lt;/h3&gt;
&lt;p&gt;Consider the polynomial \( P(x)={@simpcubic@} \)&lt;/p&gt;
&lt;p&gt;We observe that \( P( {@n1@})=0\). Then \({@n1@}\) is a root for \(P(x)=0\)&lt;/p&gt;
&lt;p&gt;By dividing {@simpcubic@} by {@x-n1@}, we obtain \({@simpquadratic[1]@}\).&lt;/p&gt;
&lt;p&gt;So, \( {@simpcubic@} = ({@x-n1@})({@simpquadratic[1]@}) \) &lt;/p&gt;
&lt;p&gt;Since  \({@simpcubic@}=0\), then \( {@x-n1=0@} \) or \( {@simpquadratic[1]=0@} \)&lt;/p&gt;
&lt;p&gt;Let solve \( {@simpquadratic[1]=0@} \). When we compute its discriminant \(\Delta=b^2-4ac\) with \(b={@bcoeff@}, a={@acoeff@}, c={@ccoeff@}\) we get, \(\Delta ={@bcoeff^2-4*acoeff*ccoeff@}&lt;0\)&lt;/p&gt;
 &lt;p&gt;So the equation \( {@simpquadratic[1]=0@} \) admits two distinct complex roots satisfying the formula; &lt;/p&gt;
&lt;p&gt; \[x=\frac{-b±\sqrt{- \Delta} {@%i@}}{2a}\] &lt;/p&gt;
&lt;p&gt; \[x=\frac{{@if bcoeff&lt;0 then "+" else "-"@}{@abs(bcoeff)@}±\sqrt{{@-(bcoeff^2-4*acoeff*ccoeff)@}} {@%i@}}{{@2*acoeff@}}\] &lt;/p&gt;
&lt;p&gt;\[x = {@z1@} \; \text{or} \;{@z2@}\]&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>n1:rand(9)+2;
 n3:rand(4)+2;
 arg1:rand([Pi/6,Pi/4,Pi/3]);
 z1:expand(n3*(cos(arg1)+i*sin(arg1)));
 z2:expand(n3*(cos(arg1)-i*sin(arg1)));
 cubequ:expand((x-n1)*(x-z1)*(x-z2));
 simpcubic:collectterms(cubequ,x);
 constcoeff:coeff(simpcubic,x,0);
 simpquadratic:divide(simpcubic,(x-n1),x);
 acoeff:coeff(simpquadratic[1],x,2);
 bcoeff:coeff(simpquadratic[1],x,1);
 ccoeff:coeff(simpquadratic[1],x,0);
 wrongz1:n3*(cos(arg1)-i*sin(arg1));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt; The cubic equation \( {@simpcubic@} = 0.\) &lt;/p&gt;
&lt;p&gt;A complex root of the cubic equation: \({@z1@}\) or  \({@z2@}\)&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>z1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>argument</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>Result</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>expandans1:expand(ans1);
/* Node 1: Checks for the correct answer */
/* Node 2: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>z1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>z2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-1-T </trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-1-F </falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>373711559</deployedseed>
    <deployedseed>817986222</deployedseed>
    <deployedseed>1331635103</deployedseed>
    <deployedseed>1009002088</deployedseed>
    <deployedseed>654432205</deployedseed>
    <deployedseed>1768193270</deployedseed>
    <deployedseed>2111310956</deployedseed>
    <deployedseed>1057980489</deployedseed>
    <deployedseed>697725931</deployedseed>
    <deployedseed>1759338217</deployedseed>
    <deployedseed>924335502</deployedseed>
    <deployedseed>1102488915</deployedseed>
    <deployedseed>2053616323</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>z1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>z2</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
