<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/The-Real-and-Complex-Number-Systems-(NS)/Polar-form-of-a-complex-number/Conversion-to-Polar-Form-Complex-Number-z-=-x-+-yi.xml">
  <question type="stack">
    <name>
      <text>Conversion to Polar Form Complex Number z = x + yi</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Write the complex number \(z={@x@}+{@y@}i\) in polar form using the argument \(\theta \) "the angle measured from the positive real axis.&lt;/p&gt;
&lt;p&gt;&lt;em&gt;Note: \({@%i @}\) can be typed as &lt;code&gt;%i&lt;/code&gt; .&lt;/em&gt;&lt;/p&gt;
&lt;p&gt;r = [[input:ans1]] and \(\theta\) = [[input:ans2]] [[validation:ans1]][[validation:ans2]]&lt;/p&gt;
&lt;p&gt;z = [[input:ans3]] [[validation:ans3]][[feedback:prt1]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked Solution&lt;/h3&gt;
&lt;p&gt;The given complex number is \({@x@}{@y@}i\)&lt;/p&gt;
&lt;p&gt;&lt;strong&gt;Step 1:&lt;/strong&gt; We calculate the modulus (magnitude) of the complex number using the formula: \(|z| = \sqrt{a^2 + b^2}\) for any complex number \(z= a+b.i \)&lt;/p&gt;
&lt;p&gt;So, for our complex number the modulus is: \(\displaystyle r=\sqrt{({@x@}^2)+({@y@}^2)}\)&lt;/p&gt;
&lt;p&gt;Modulus r={@r@}&lt;/p&gt;
&lt;p&gt;&lt;/p&gt;
&lt;p&gt;&lt;strong&gt;Step 2:&lt;/strong&gt; The argument \( \theta \) of a complex number can be found using the formula:&lt;/p&gt;
&lt;p&gt;\( \displaystyle \theta =tan^-1 \left( \frac{b}{a}\right) \implies tan^-1 \left( \frac{{@y@}}{{@x@}}\right) \) (Note: The angle needs to be in radians)&lt;/p&gt;
&lt;p&gt;\(\theta ={@arg@}\)&lt;/p&gt;
&lt;p&gt;&lt;/p&gt;
&lt;p&gt;&lt;strong&gt;Step 3:&lt;/strong&gt; The polar form of the complex number is given by \(z=r(\cos\theta+i\,\sin\theta)\)&lt;/p&gt;
&lt;p&gt;Substituting the values of \(r\) and \(\theta\) into the formula: \( z = {@r@}\left(\cos{@arg@}+i\, \sin {@arg@}\right)\)&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>a: 1+rand(10);
b: a*sqrt(3);
z: -1*a-b*%i;
x:realpart(z);
y:imagpart(z);
r:  cabs(z);
az: atan(y/x);
arg: %pi +  carg(z);
pz: r*%e^(%i*arg);
ta: pz;
simp:false;
ta2: r*(cos(arg)+%i*sin(arg));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;The complex number \(z={@x@}{@y@}i\) in polar form using the argument.&lt;/p&gt;
&lt;p&gt;&lt;strong&gt;Give your solution in \(\pi\) radians, or in the form \(z=re^{i\theta}\).&lt;/strong&gt;&lt;/p&gt;
&lt;p&gt;\(r ={@r@}\) and \(\theta ={@arg@}\)&lt;/p&gt;
&lt;p&gt;\(z ={@ta@}\) &lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;p&gt;Correct answer, well done.&lt;/p&gt;</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;Your answer is partially correct.&lt;/p&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;p&gt;Incorrect answer.&lt;/p&gt;</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>r</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>arg</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans: ans1*%e^(%i*ans2);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>r</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Check that you have calculated the value of \(r\) correctly.&lt;/p&gt;</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>arg</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Check that you have correctly measured the angle from the positive real axis.&lt;/p&gt;</text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>sans</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;Your values for r and \(\theta\) are incorrect, but you have correctly written the polar form using these values.&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Your values for r and \(\theta\) are correct, but you have not correctly written the polar form for the number. Check your syntax!&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>800765924</deployedseed>
    <deployedseed>1560814259</deployedseed>
    <deployedseed>2110675066</deployedseed>
    <deployedseed>973220628</deployedseed>
    <deployedseed>334938587</deployedseed>
    <deployedseed>1780174110</deployedseed>
    <deployedseed>649209433</deployedseed>
    <deployedseed>1133843212</deployedseed>
    <deployedseed>1775401576</deployedseed>
    <deployedseed>50388741</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>r</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>arg</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
