<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Propositional-Logic-and-Set-Theory-(NS-and-SS)/Definition-and-examples-of-propositions/Determine-which-statements-are-true.xml">
  <question type="stack">
    <name>
      <text>Determine which statements are true</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt; Select all of the statements below that are true.&lt;/p&gt;
&lt;p&gt;[[input:ans1]] [[validation:ans1]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked solution&lt;/h3&gt;

[[ if test="member(1, ta_index)" ]]
&lt;p&gt; "Horses are mammals and frogs are amphibians."&lt;/p&gt;
&lt;p&gt;This statement is true because both of the comprising statements are true.&lt;/p&gt;
{@if not(ta_largest=1) then "&lt;br&gt;" else ""@}
[[/ if ]]

[[ if test="member(2, ta_index)" ]]
&lt;p&gt;"The sum of the interior angles of any triangle is always 180 degrees."&lt;/p&gt;
&lt;p&gt;The Triangle Sum Theorem states that the sum of the interior angles of a triangle is always equal to a straight angle, which measures 180 degrees thus the statement is true.&lt;/p&gt;
{@if not(ta_largest=1) then "&lt;br&gt;" else ""@}
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
&lt;p&gt;"By definition, in Euclidean geometry, parallel lines are lines in a plane that do not intersect"&lt;/p&gt;
&lt;p&gt;Two parallel lines will never intersect because they maintain a constant distance from each other at every point hence the statement is true.&lt;/p&gt;
{@if not(ta_largest=1) then "&lt;br&gt;" else ""@}
[[/ if ]]

[[ if test="member(4, ta_index)" ]]
&lt;p&gt;"\(2 + 4 = 8\)"&lt;/p&gt;
&lt;p&gt; Here, we are claiming that \(2 + 4 = 8\) which is &lt;b&gt;False&lt;/b&gt;. 
&lt;br&gt;Therefore, this is a false statement.&lt;/p&gt;
{@if not(ta_largest=1) then "&lt;br&gt;" else ""@}
[[/ if ]]

[[ if test="member(5, ta_index)" ]]
&lt;p&gt;"Division by zero is defined "&lt;/p&gt;
&lt;p&gt;The statement is false because in mathematics, division by zero is considered undefined because it leads to various contradictions and inconsistencies. When we divide a number by a very small positive number, the result becomes extremely large. As the divisor approaches zero, the quotient tends to infinity, which is not a finite and well-defined value. On the other hand, if the divisor approaches zero from the negative side, the quotient tends to negative infinity.&lt;/p&gt;
{@if not(ta_largest=1) then "&lt;br&gt;" else ""@}
[[/ if ]]

[[ if test="member(6, ta_index)" ]]
&lt;p&gt;"The square root of a negative number is a real number."&lt;/p&gt;
&lt;p&gt;The statement is false because the square root of a negative number is not a real number. In the real number system, there is no real number that, when squared, gives a negative result. This is because the square of any real number is always non-negative.&lt;/p&gt;
{@if not(ta_largest=1) then "&lt;br&gt;" else ""@}
[[/ if ]]</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>ta_true: rand_selection([
  [1, true, "Horses are mammals and frogs are amphibians."],
  [2, true, "The sum of the interior angles of any triangle in Euclidean geometry is always 180 degrees.."],
  [3, true, "By definition, in Euclidean geometry, parallel lines are lines in a plane that do not intersect"]
], 1);

ta_false: rand_selection([
  [4, false, "\\(2 + 4 = 8\\)"],
  [5, false, "Division by zero is defined"],
  [6, false, "The square root of a negative number is a real number."]
], 2);

ta:random_permutation(append(ta_true,ta_false));
ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;True statement: {@mcq_correct(ta)[1]@}&lt;/p&gt;
&lt;p&gt;False statements: {@mcq_incorrect(ta)@}&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt; Your answer is partially correct.&lt;/span&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>415442381</deployedseed>
    <deployedseed>60683925</deployedseed>
    <deployedseed>1487461689</deployedseed>
    <deployedseed>499706999</deployedseed>
    <deployedseed>1284588220</deployedseed>
    <deployedseed>957474342</deployedseed>
    <deployedseed>1273001315</deployedseed>
    <deployedseed>31931398</deployedseed>
    <deployedseed>1130112719</deployedseed>
    <deployedseed>620883387</deployedseed>
    <deployedseed>1993736347</deployedseed>
    <deployedseed>1062169906</deployedseed>
    <deployedseed>1883997875</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
