<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Functions-(NS-and-SS)/Real-valued-functions-and-their-properties/Domain-and-Range-of-a-Function_3.xml">
  <question type="stack">
    <name>
      <text>Domain and Range of a Function</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Find the domain and range of \(f(x)={@fx@}\).&lt;/p&gt;
&lt;p&gt;Give your answer as interval of the real numbers.&lt;/p&gt;

&lt;p&gt;&lt;em&gt;Note: To enter intervals of you real line, you should write the relevant text in the table below.&lt;/em&gt;&lt;/p&gt;

&lt;table width="250" cellspacing="1" cellpadding="0" border="1" style="margin-left: 30px"&gt;
  &lt;tr&gt;
    &lt;th style="text-align: center;"&gt;Intervals&lt;/th&gt;
    &lt;th style="text-align: center;"&gt;Input Syntax&lt;/th&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td style="text-align: center;"&gt;\( [a,b] \)&lt;/td&gt;
    &lt;td style="text-align: center;"&gt;&lt;code&gt;cc(a,b)&lt;/code&gt;&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td style="text-align: center;"&gt;\( (a,b) \)&lt;/td&gt;
    &lt;td style="text-align: center;"&gt;&lt;code&gt;oo(a,b)&lt;/code&gt;&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td style="text-align: center;"&gt;\( [a,b) \)&lt;/td&gt;
    &lt;td style="text-align: center;"&gt;&lt;code&gt;co(a,b)&lt;/code&gt;&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td style="text-align: center;"&gt;\( (a,b] \)&lt;/td&gt;
    &lt;td style="text-align: center;"&gt;&lt;code&gt;oc(a,b)&lt;/code&gt;&lt;/td&gt;
  &lt;/tr&gt;
 &lt;tr&gt;
    &lt;td style="text-align: center;"&gt;\( (a, \infty) \)&lt;/td&gt;
    &lt;td style="text-align: center;"&gt;&lt;code&gt;oo(a,inf)&lt;/code&gt;&lt;/td&gt;
  &lt;/tr&gt;
 &lt;tr&gt;
    &lt;td style="text-align: center;"&gt;\( (-\infty, b] \)&lt;/td&gt;
    &lt;td style="text-align: center;"&gt;&lt;code&gt;oc(-inf,b)&lt;/code&gt;&lt;/td&gt;
  &lt;/tr&gt;
 &lt;tr&gt;
    &lt;td style="text-align: center;"&gt;\( (-\infty, \infty) \)&lt;/td&gt;
    &lt;td style="text-align: center;"&gt;&lt;code&gt;oo(-inf,inf)&lt;/code&gt;&lt;/td&gt;
  &lt;/tr&gt;
&lt;/table&gt;

&lt;br&gt;

&lt;p&gt;(a) \(Dom(f)=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]&lt;/p&gt;
&lt;p&gt;(b) \(Range(f)=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]&lt;br&gt;&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked Solution&lt;/h3&gt;

&lt;p&gt;(a) To determine the domain of the function, we need to identify any restrictions on the values of \(x\) that would make the expression undefined.&lt;/p&gt;

&lt;p&gt;Looking at the function expression \(f(x)={@fx@}\), there are no specific restrictions on the input value \( x \).&lt;/p&gt;
&lt;p&gt;The function \( x \mapsto x^2 - {@bb@} \) is a polynomial function, so defined for all real numbers \( x \). Moreover, the absolute value function \( \lvert \cdot \rvert \) is defined for all real numbers.&lt;/p&gt; 
&lt;p&gt;Therefore, the domain of \(f(x)\) is all real numbers, which can be denoted as,&lt;/p&gt;
&lt;p style="margin-left: 30px"&gt;\( Dom(f) = \mathbb{R} = {@ta1@}\)&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(b) Next, we determine the range of the function. 
&lt;br&gt;The range represents the set of all possible values that \(f(x)\) can attain as \(x\) varies.&lt;/p&gt;

&lt;p&gt;We know that, for all \(x \in \mathbb{R}\)
&lt;p style="margin-left: 30px"&gt;\(\begin{align}
&amp; {@abs(abs(x^2-bb))@} \geq 0 \\
&amp; \Rightarrow {@cc*abs(abs(x^2-bb))@} \geq 0 \\
&amp; \Rightarrow {@cc*abs(abs(x^2-bb))+aa@} \geq {@aa@} \\
\end{align}\)&lt;/p&gt;

&lt;p&gt;Hence, for all \(x \in \mathbb{R}\), \(f(x) \geq {@aa@}\), so,&lt;/p&gt;
&lt;p style="margin-left: 30px"&gt;\(Range(f) = \{ y : y \in \mathbb{R}, y \geq {@aa@} \} = {@ta2@} \)&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>pm1:rand_with_prohib(-1,1,[0]);
aa:pm1*(rand(9)+1);
bb:rand([1,4,9,16,25]);
cc:rand_with_prohib(1,10,[aa]);
fx:aa+cc*abs(x^2-bb);
ta1:oo(-inf,inf);
ta2:co(aa,inf);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;Given a function \(f(x)={@fx@}\).&lt;/p&gt;
&lt;p&gt;Domain \(={@ta1@}\) &lt;/p&gt;
&lt;p&gt;Range  \(={@ta2@}\)&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;The absolute value of any real number can be calculated, so there are no restrictions of \(x\) to be considered. Therefore the domain of \(f(x)\) is \(\mathbb{R}\).&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;We know that \( \forall x \in \mathbb{R}, {@cc*abs(abs(x^2-bb))@} \geq 0\). &lt;/p&gt;
&lt;p&gt; We then need to add {@aa@} to the absolute value inequality, making the minimal value of \(f(x)\) to be \({@aa@}\). &lt;/p&gt;
&lt;p&gt;Therefore, the Range of the function is \(y \geq {@aa@}\).&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2123903877</deployedseed>
    <deployedseed>1844287144</deployedseed>
    <deployedseed>1809629696</deployedseed>
    <deployedseed>1900403505</deployedseed>
    <deployedseed>1172961410</deployedseed>
    <deployedseed>1749634100</deployedseed>
    <deployedseed>974586795</deployedseed>
    <deployedseed>1301300332</deployedseed>
    <deployedseed>1275039800</deployedseed>
    <deployedseed>1752222885</deployedseed>
    <deployedseed>368304658</deployedseed>
    <deployedseed>664271737</deployedseed>
    <deployedseed>1956377741</deployedseed>
    <deployedseed>1944262633</deployedseed>
    <deployedseed>793549167</deployedseed>
    <deployedseed>391123462</deployedseed>
    <deployedseed>1798646848</deployedseed>
    <deployedseed>276826719</deployedseed>
    <deployedseed>880085380</deployedseed>
    <deployedseed>1695806276</deployedseed>
    <deployedseed>1915679204</deployedseed>
    <deployedseed>920709668</deployedseed>
    <deployedseed>498558570</deployedseed>
    <deployedseed>1609004559</deployedseed>
    <deployedseed>743716200</deployedseed>
    <deployedseed>154028772</deployedseed>
    <deployedseed>1950905518</deployedseed>
    <deployedseed>587993355</deployedseed>
    <deployedseed>410882256</deployedseed>
    <deployedseed>1566589612</deployedseed>
    <deployedseed>1162667797</deployedseed>
    <deployedseed>2049871267</deployedseed>
    <deployedseed>1959934461</deployedseed>
    <deployedseed>1323821639</deployedseed>
    <deployedseed>1088915163</deployedseed>
    <deployedseed>1798418753</deployedseed>
    <deployedseed>188731367</deployedseed>
    <deployedseed>756144633</deployedseed>
    <deployedseed>1827825423</deployedseed>
    <deployedseed>899463819</deployedseed>
    <deployedseed>1601119327</deployedseed>
    <deployedseed>375858142</deployedseed>
    <deployedseed>1887672446</deployedseed>
    <deployedseed>1186200122</deployedseed>
    <deployedseed>1794633661</deployedseed>
    <deployedseed>2112737694</deployedseed>
    <deployedseed>2127151944</deployedseed>
    <deployedseed>2001849986</deployedseed>
    <deployedseed>1758034608</deployedseed>
    <deployedseed>1756204791</deployedseed>
    <deployedseed>1877647598</deployedseed>
    <deployedseed>168355032</deployedseed>
    <deployedseed>2141987717</deployedseed>
    <deployedseed>884287486</deployedseed>
    <deployedseed>451294862</deployedseed>
    <deployedseed>1854078574</deployedseed>
    <deployedseed>887277348</deployedseed>
    <deployedseed>1162698165</deployedseed>
    <deployedseed>834454707</deployedseed>
    <deployedseed>1155646494</deployedseed>
    <deployedseed>1584423395</deployedseed>
    <deployedseed>1826799018</deployedseed>
    <deployedseed>1783305999</deployedseed>
    <deployedseed>393758383</deployedseed>
    <deployedseed>532847482</deployedseed>
    <deployedseed>1089998821</deployedseed>
    <deployedseed>101215325</deployedseed>
    <deployedseed>1720667909</deployedseed>
    <deployedseed>1788826396</deployedseed>
    <deployedseed>750306327</deployedseed>
    <deployedseed>847561207</deployedseed>
    <deployedseed>984201392</deployedseed>
    <deployedseed>722883420</deployedseed>
    <deployedseed>569873556</deployedseed>
    <deployedseed>844011050</deployedseed>
    <deployedseed>1977884</deployedseed>
    <deployedseed>902315215</deployedseed>
    <deployedseed>827116324</deployedseed>
    <deployedseed>1096431946</deployedseed>
    <deployedseed>2142469332</deployedseed>
    <deployedseed>773393513</deployedseed>
    <deployedseed>1882676086</deployedseed>
    <deployedseed>348069197</deployedseed>
    <deployedseed>20449911</deployedseed>
    <deployedseed>529583219</deployedseed>
    <deployedseed>1368230260</deployedseed>
    <deployedseed>1957610760</deployedseed>
    <deployedseed>1977049876</deployedseed>
    <deployedseed>1731506588</deployedseed>
    <deployedseed>1243389599</deployedseed>
    <deployedseed>1995584202</deployedseed>
    <deployedseed>1593651757</deployedseed>
    <deployedseed>1814126743</deployedseed>
    <deployedseed>1587740511</deployedseed>
    <deployedseed>33127088</deployedseed>
    <deployedseed>1562205490</deployedseed>
    <deployedseed>485122213</deployedseed>
    <deployedseed>2012665597</deployedseed>
    <deployedseed>706098139</deployedseed>
    <deployedseed>264811151</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
