<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Functions-(NS-and-SS)/Types-of-functions-and-inverse-of-a-function/Finding-inverse-of-a-function-g(x).xml">
  <question type="stack">
    <name>
      <text>Finding inverse of a function g(x)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text>&lt;p&gt;If \(g(x)={@const_1@}+{@coeff_x*x@}+e^x\), find \(g^{-1}({@const_1+1@})\)&lt;/p&gt;
&lt;p&gt; \(g^{-1}({@const_1+1@})=\) [[input:ans1]][[validation:ans1]][[feedback:prt1]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text>&lt;h3&gt;Worked Solution&lt;/h3&gt;
&lt;p&gt;In this case we cannot find \( g^{-1}(x), \) however we can use some elements of the domain of \(g\) to give us the image \({@const_1+1@},\) &lt;/p&gt;
&lt;p&gt;\(g(-2)={@const_1@}+{@coeff_x@}(-2)+e^{-2}={@const_1+coeff_x*(-2)@}+\frac{1}{e^{2}}\)&lt;/p&gt;
&lt;p&gt;\(g(-1)={@const_1@}+{@coeff_x@}(-1)+e^{-1}={@const_1+coeff_x*(-1)@}+\frac{1}{e}\)&lt;/p&gt;
&lt;p&gt;\(g(0)={@const_1@}+0+e^0={@const_1+coeff_x*(0)@}+1={@const_1+1@}\)&lt;/p&gt;
&lt;p&gt;Thus, \(g^{-1}({@const_1+1@})=0.\)&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>const_1:rand(7)+1;
coeff_x:rand(5)+1;
ta:0;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt; \(g(x)={@const_1@}+{@coeff_x*x@}+e^x\)&lt;/p&gt;
&lt;p&gt; \(g^{-1}({@const_1+1@})={@ta@}\)&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="html">
      <text>&lt;p id="author"&gt;Nixson Kiplagat&lt;/p&gt;
&lt;p id="concept"&gt;https://docs.google.com/document/d/1A5oyYMygAgYhjif6EiN3EjMlvtwS4gkLIbsD9gR0P-I/edit&lt;/p&gt;
&lt;p id="reviewer"&gt;Emmaculate Odhiambo&lt;/p&gt;
&lt;p id="reviewer"&gt;Santiago Borio&lt;/p&gt;
&lt;p id="description"&gt;Given a function g(x), needs to find its inverse&lt;/p&gt;</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/
/*Node 2:  Finding the value of the function rather than the inverse*/
x_value: const_1 + 1;
g_x:const_1+coeff_x*x+e^x;
mistake2_1:ev(g_x, x = x_value);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description> Finding the value of the function rather than the inverse</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake2_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;You seem to have found the value of the function rather than the inverse which is not true. Check the worked Solution below&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>971117956</deployedseed>
    <deployedseed>1205573946</deployedseed>
    <deployedseed>1104320778</deployedseed>
    <deployedseed>622738940</deployedseed>
    <deployedseed>1566507527</deployedseed>
    <deployedseed>869034020</deployedseed>
    <deployedseed>290144473</deployedseed>
    <deployedseed>1275379484</deployedseed>
    <deployedseed>629231098</deployedseed>
    <deployedseed>1472538034</deployedseed>
    <deployedseed>1606903226</deployedseed>
    <deployedseed>609400881</deployedseed>
    <deployedseed>846809272</deployedseed>
    <deployedseed>982166083</deployedseed>
    <deployedseed>487039082</deployedseed>
    <deployedseed>177707368</deployedseed>
    <deployedseed>1363529812</deployedseed>
    <deployedseed>847181066</deployedseed>
    <deployedseed>1992158213</deployedseed>
    <deployedseed>2105696292</deployedseed>
    <deployedseed>618775502</deployedseed>
    <deployedseed>856915581</deployedseed>
    <deployedseed>715319142</deployedseed>
    <deployedseed>804887948</deployedseed>
    <deployedseed>885468190</deployedseed>
    <deployedseed>1766949620</deployedseed>
    <deployedseed>924294338</deployedseed>
    <deployedseed>1529365401</deployedseed>
    <deployedseed>527517734</deployedseed>
    <deployedseed>92442759</deployedseed>
    <deployedseed>1760307292</deployedseed>
    <deployedseed>1386641429</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
