<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Propositional-Logic-and-Set-Theory-(NS-and-SS)/Logical-connectives/Give-statement-forms-of-the-axioms-of-an-enigma-and-answer-to-an.xml">
<!-- question: 222435  -->
  <question type="stack">
    <name>
      <text>Give statement forms of the axioms of an enigma and answer to an enigma. (narrow formatting)</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Consider the following enigma problem, and give statement forms of its axioms. Then fill in the truth table, and answer to the enigma: &lt;em&gt;"Who order a dessert and not"&lt;/em&gt;.&lt;/p&gt;
&lt;hr /&gt;
&lt;p&gt;&lt;strong&gt;Problem:&lt;/strong&gt;Three Colleagues, Albert, Bernard and Charles, have lunch together every working day. &lt;br /&gt;The following axioms are true:&lt;/p&gt;
&lt;ol&gt;
&lt;li&gt;If Albert orders a dessert, Bernard orders one too.&lt;/li&gt;
&lt;li&gt;Every day, either Bernard or Charles, but not both, orders a dessert.&lt;/li&gt;
&lt;li&gt;Albert or Charles, or both, order dessert daily.&lt;/li&gt;
&lt;li&gt;If Charles orders dessert, Albert does the same.&lt;/li&gt;
&lt;/ol&gt;&lt;hr /&gt;
&lt;p&gt;&lt;strong&gt;Statement forms of axioms&lt;/strong&gt; &lt;br /&gt;Give the statement form of each axiom given in the problem. &lt;br /&gt;Let a: "Albert orders a dessert",      b: "Bernard orders a dessert", and c: "Charles orders a dessert". &lt;br /&gt;&lt;span style="font-size: 0.9375rem;"&gt;&lt;span style="font-size: 0.9375rem;"&gt;1. "If Albert orders a dessert, Bernard orders one too": [[input:ans34]] [[validation:ans34]]  [[feedback:prt10]] &lt;br /&gt;&lt;span style="font-size: 0.9375rem;"&gt;2. "Every day, either Bernard or Charles, but not both, orders a dessert": [[input:ans35]] [[validation:ans35]]  [[feedback:prt11]] &lt;br /&gt;&lt;span style="font-size: 0.9375rem;"&gt;3. "Albert or Charles, or both, order dessert daily": [[input:ans36]] [[validation:ans36]]  [[feedback:prt12]] &lt;br /&gt;&lt;span style="font-size: 0.9375rem;"&gt;4. "If Charles orders dessert, Albert does the same": [[input:ans37]] [[validation:ans37]]  [[feedback:prt13]] &lt;/p&gt;
&lt;center&gt;
&lt;p&gt;&lt;strong&gt;Fill in truth table&lt;/strong&gt;&lt;/p&gt;

&lt;table style="border-collapse: collapse; width: 80%; text-align: center; font-size: 14px; margin-left: 0; margin-right: auto;" border="1"&gt;
  &lt;tr&gt;
    &lt;th style="width: 7%;"&gt;\(a\)&lt;/th&gt;
    &lt;th style="width: 7%;"&gt;\(b\)&lt;/th&gt;
    &lt;th style="width: 7%;"&gt;\(c\)&lt;/th&gt;
    &lt;th style="width: 12%;"&gt;\(a \implies b\)&lt;/th&gt;
    &lt;th style="width: 12%;"&gt;\((b \land \neg c) \lor (\neg b \land c)\)&lt;/th&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans1]]&lt;/td&gt;
    &lt;td&gt;[[input:ans9]]&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans2]]&lt;/td&gt;
    &lt;td&gt;[[input:ans10]]&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans3]]&lt;/td&gt;
    &lt;td&gt;[[input:ans11]]&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans4]]&lt;/td&gt;
    &lt;td&gt;[[input:ans12]]&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans5]]&lt;/td&gt;
    &lt;td&gt;[[input:ans13]]&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans6]]&lt;/td&gt;
    &lt;td&gt;[[input:ans14]]&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans7]]&lt;/td&gt;
    &lt;td&gt;[[input:ans15]]&lt;/td&gt;
  &lt;/tr&gt;
  &lt;tr&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans8]]&lt;/td&gt;
    &lt;td&gt;[[input:ans16]]&lt;/td&gt;
  &lt;/tr&gt;
&lt;/table&gt;

&lt;table style="border-collapse: collapse; width: 80%; text-align: center; font-size: 14px; margin-left: 0; margin-right: auto;" border="1"&gt;
  &lt;tr&gt;
    &lt;th style="width: 7%;"&gt;\(a\)&lt;/th&gt;
    &lt;th style="width: 7%;"&gt;\(b\)&lt;/th&gt;
    &lt;th style="width: 7%;"&gt;\(c\)&lt;/th&gt;
    &lt;th style="width: 12%;"&gt;\(a \lor c\)&lt;/th&gt;
    &lt;th style="width: 12%;"&gt;\(c \implies a\)&lt;/th&gt;
    &lt;th style="width: 11%;"&gt;Feedback&lt;/th&gt;
  &lt;/tr&gt;

  &lt;tr&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans17]]&lt;/td&gt;
    &lt;td&gt;[[input:ans25]]&lt;/td&gt;
    &lt;td&gt;[[feedback:prt1]]&lt;/td&gt;
  &lt;/tr&gt;

  &lt;tr&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans18]]&lt;/td&gt;
    &lt;td&gt;[[input:ans26]]&lt;/td&gt;
    &lt;td&gt;[[feedback:prt2]]&lt;/td&gt;
  &lt;/tr&gt;

  &lt;tr&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans19]]&lt;/td&gt;
    &lt;td&gt;[[input:ans27]]&lt;/td&gt;
    &lt;td&gt;[[feedback:prt3]]&lt;/td&gt;
  &lt;/tr&gt;

  &lt;tr&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans20]]&lt;/td&gt;
    &lt;td&gt;[[input:ans28]]&lt;/td&gt;
    &lt;td&gt;[[feedback:prt4]]&lt;/td&gt;
  &lt;/tr&gt;

  &lt;tr&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans21]]&lt;/td&gt;
    &lt;td&gt;[[input:ans29]]&lt;/td&gt;
    &lt;td&gt;[[feedback:prt5]]&lt;/td&gt;
  &lt;/tr&gt;

  &lt;tr&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[1]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans22]]&lt;/td&gt;
    &lt;td&gt;[[input:ans30]]&lt;/td&gt;
    &lt;td&gt;[[feedback:prt6]]&lt;/td&gt;
  &lt;/tr&gt;

  &lt;tr&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[1]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans23]]&lt;/td&gt;
    &lt;td&gt;[[input:ans31]]&lt;/td&gt;
    &lt;td&gt;[[feedback:prt7]]&lt;/td&gt;
  &lt;/tr&gt;

  &lt;tr&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@q[2]@}\)&lt;/td&gt;
    &lt;td&gt;\({@p[2]@}\)&lt;/td&gt;
    &lt;td&gt;[[input:ans24]]&lt;/td&gt;
    &lt;td&gt;[[input:ans32]]&lt;/td&gt;
    &lt;td&gt;[[feedback:prt8]]&lt;/td&gt;
  &lt;/tr&gt;
&lt;/table&gt;

&lt;/center&gt;
&lt;p&gt;&lt;/p&gt;
&lt;p&gt;&lt;strong&gt;Answer to the enigma:&lt;/strong&gt; Who orders a dessert?&lt;/p&gt;
&lt;p&gt;&lt;em&gt;Note: Hint: Observe from the truth table the interpretation that makes all four statement forms true. &lt;/em&gt;&lt;/p&gt;
&lt;p&gt; [[input:ans33]] [[validation:ans33]]  [[feedback:prt9]]&lt;/p&gt;
&lt;p&gt;[[validation:ans1]] [[validation:ans2]] [[validation:ans3]] [[validation:ans4]] [[validation:ans5]] [[validation:ans6]] [[validation:ans7]] [[validation:ans8]] [[validation:ans17]] [[validation:ans18]] [[validation:ans19]] [[validation:ans20]]&lt;/p&gt;
&lt;p&gt;[[validation:ans9]] [[validation:ans10]] [[validation:ans11]] [[validation:ans12]] [[validation:ans13]] [[validation:ans14]] [[validation:ans15]] [[validation:ans16]] [[validation:ans21]] [[validation:ans22]] [[validation:ans23]] [[validation:ans24]]&lt;/p&gt;
&lt;p&gt;[[validation:ans25]] [[validation:ans26]] [[validation:ans27]] [[validation:ans28]][[validation:ans29]] [[validation:ans30]] [[validation:ans31]] [[validation:ans32]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked solution&lt;/h3&gt;
&lt;p&gt;&lt;strong&gt;Statement forms of axioms:&lt;/strong&gt; &lt;br /&gt;Give the statement form of each axiom given in the problem. &lt;br /&gt;Let a: "Albert orders a dessert", b: "Bernard orders a dessert", and c: "Charles orders a dessert". &lt;br /&gt;1. "If Albert orders a dessert, Bernard orders one too": a implies b &lt;br /&gt;2. "Every day, either Bernard or Charles, but not both, orders a dessert": (b and not c) or (not b and c) &lt;br /&gt;3. "Albert or Charles, or both, order dessert daily": a or c &lt;br /&gt;4. "If Charles orders dessert, Albert does the same": c implies a&lt;/p&gt;
&lt;hr /&gt;
&lt;p&gt;&lt;strong&gt;Fill in truth table:&lt;/strong&gt; Here are the truth tables of each of these statements &lt;/p&gt;
&lt;p&gt;Truth table of \(a \implies b\): {@truth_table(ta1)@} &lt;/p&gt;
&lt;p&gt;Truth table of ( \( b \land \neg c\) ) \(\lor\)  ( \(\neg b \land c\) ): {@truth_table(P1)@}&lt;/p&gt;
&lt;p&gt;Truth table of \( a \lor c \): {@truth_table(ta3)@} &lt;/p&gt;
&lt;p&gt;Truth table of \( c \implies a \): {@truth_table(ta4)@}&lt;/p&gt;
&lt;hr /&gt;
&lt;p&gt;&lt;strong&gt;Answer to the enigma:&lt;/strong&gt; Who order a dessert and who does not order a dessert? &lt;br /&gt;We deduction that the interpretation that makes all four statement forms true is the second line from the truth table. &lt;br /&gt;One can see that only "a" and "b" are true and "c" is false. Therefore, Albert and Bernard order dessert but not Charles.&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>9</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>p: ( ["True","False"] );
q: ( ["True","False"] );
P1: (b nounand not c) nounor (not b nounand c);
ta1: a implies b;
ta2: (b nounand not c) nounor (not b nounand c);
ta3: a nounor c;
ta4: c implies a;
tas: [[true, false, true, true], [true, true, true, true], [false, true, true, true], [false, false, true, true], [true, false, true, false], [true, true, false, true], [true, true, true, false], [true, false, false, true]];
ta: random_permutation([["Bernard and Charles order dessert but not Albert.", false],["Albert and Bernard order dessert but not Charles.", true],["Albert and Charles order dessert but not Bernard.", false]]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;Note: There is no randomization in this question&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;p&gt;Correct answer, well done.&lt;/p&gt;</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;&lt;span style="font-size: 1.5em; color: orange;"&gt; Your answer is partially correct.&lt;/span&gt;&lt;/p&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;p&gt;Incorrect answer.&lt;/p&gt;</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>boolean</type>
      <tans>tas[1][1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>boolean</type>
      <tans>tas[2][2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>boolean</type>
      <tans>tas[3][2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>boolean</type>
      <tans>tas[4][2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13</name>
      <type>boolean</type>
      <tans>tas[5][2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans14</name>
      <type>boolean</type>
      <tans>tas[6][2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans15</name>
      <type>boolean</type>
      <tans>tas[7][2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans16</name>
      <type>boolean</type>
      <tans>tas[8][2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans17</name>
      <type>boolean</type>
      <tans>tas[1][3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans18</name>
      <type>boolean</type>
      <tans>tas[2][3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans19</name>
      <type>boolean</type>
      <tans>tas[3][3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>boolean</type>
      <tans>tas[2][1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans20</name>
      <type>boolean</type>
      <tans>tas[4][3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans21</name>
      <type>boolean</type>
      <tans>tas[5][3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans22</name>
      <type>boolean</type>
      <tans>tas[6][3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans23</name>
      <type>boolean</type>
      <tans>tas[7][3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans24</name>
      <type>boolean</type>
      <tans>tas[8][3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans25</name>
      <type>boolean</type>
      <tans>tas[1][4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans26</name>
      <type>boolean</type>
      <tans>tas[2][4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans27</name>
      <type>boolean</type>
      <tans>tas[3][4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans28</name>
      <type>boolean</type>
      <tans>tas[4][4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans29</name>
      <type>boolean</type>
      <tans>tas[5][4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>boolean</type>
      <tans>tas[3][1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans30</name>
      <type>boolean</type>
      <tans>tas[6][4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans31</name>
      <type>boolean</type>
      <tans>tas[7][4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans32</name>
      <type>boolean</type>
      <tans>tas[8][4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans33</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans34</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans35</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans36</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans37</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>boolean</type>
      <tans>tas[4][1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>boolean</type>
      <tans>tas[5][1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>boolean</type>
      <tans>tas[6][1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>boolean</type>
      <tans>tas[7][1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>boolean</type>
      <tans>tas[8][1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>boolean</type>
      <tans>tas[1][2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1, ans9, ans17, ans25]</sans>
        <tans>tas[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>ans34</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>ans35</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt12</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>ans36</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt12-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt13</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>ans37</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt13-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt13-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans2, ans10, ans18, ans26]</sans>
        <tans>tas[2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans3, ans11, ans19, ans27]</sans>
        <tans>tas[3]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans4, ans12, ans20, ans28]</sans>
        <tans>tas[4]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans5, ans13, ans21, ans29]</sans>
        <tans>tas[5]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans6, ans14, ans22, ans30]</sans>
        <tans>tas[6]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans7, ans15, ans23, ans31]</sans>
        <tans>tas[7]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>[ans8, ans16, ans24, ans32]</sans>
        <tans>tas[8]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans33</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1765782283</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>tas[1][1]</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>tas[2][2]</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>tas[3][2]</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>tas[4][2]</value>
      </testinput>
      <testinput>
        <name>ans13</name>
        <value>tas[5][2]</value>
      </testinput>
      <testinput>
        <name>ans14</name>
        <value>tas[6][2]</value>
      </testinput>
      <testinput>
        <name>ans15</name>
        <value>tas[7][2]</value>
      </testinput>
      <testinput>
        <name>ans16</name>
        <value>tas[8][2]</value>
      </testinput>
      <testinput>
        <name>ans17</name>
        <value>tas[1][3]</value>
      </testinput>
      <testinput>
        <name>ans18</name>
        <value>tas[2][3]</value>
      </testinput>
      <testinput>
        <name>ans19</name>
        <value>tas[3][3]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>tas[2][1]</value>
      </testinput>
      <testinput>
        <name>ans20</name>
        <value>tas[4][3]</value>
      </testinput>
      <testinput>
        <name>ans21</name>
        <value>tas[5][3]</value>
      </testinput>
      <testinput>
        <name>ans22</name>
        <value>tas[6][3]</value>
      </testinput>
      <testinput>
        <name>ans23</name>
        <value>tas[7][3]</value>
      </testinput>
      <testinput>
        <name>ans24</name>
        <value>tas[8][3]</value>
      </testinput>
      <testinput>
        <name>ans25</name>
        <value>tas[1][4]</value>
      </testinput>
      <testinput>
        <name>ans26</name>
        <value>tas[2][4]</value>
      </testinput>
      <testinput>
        <name>ans27</name>
        <value>tas[3][4]</value>
      </testinput>
      <testinput>
        <name>ans28</name>
        <value>tas[4][4]</value>
      </testinput>
      <testinput>
        <name>ans29</name>
        <value>tas[5][4]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>tas[3][1]</value>
      </testinput>
      <testinput>
        <name>ans30</name>
        <value>tas[6][4]</value>
      </testinput>
      <testinput>
        <name>ans31</name>
        <value>tas[7][4]</value>
      </testinput>
      <testinput>
        <name>ans32</name>
        <value>tas[8][4]</value>
      </testinput>
      <testinput>
        <name>ans33</name>
        <value>mcq_correct(ta)</value>
      </testinput>
      <testinput>
        <name>ans34</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans35</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans36</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans37</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>tas[4][1]</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>tas[5][1]</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>tas[6][1]</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>tas[7][1]</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>tas[8][1]</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>tas[1][2]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt12-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt13</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt13-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>
            
