<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Analytic-Geometry-(NS)/Analysis-of-the-general-second-degree-equation/Graph-Analysis-Hyperbola-and-Intersecting-Lines-with-B-0.xml">
  <question type="stack">
    <name>
      <text>Graph Analysis: Hyperbola and Intersecting Lines with B &gt; 0</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;A conic in the two-dimensional plane is the solution set of an equation of the form&lt;/p&gt;
&lt;p&gt;\(\displaystyle ax^{2}+bxy+cy^{2}+dx+ey+f=0\)&lt;/p&gt;
&lt;p&gt;for real parameters \(\displaystyle a,b,c,d,e\) and \(f,\) where at least one of the \(a,b\) and \(c\) is not zero.&lt;/p&gt;
&lt;p&gt;If \(\displaystyle b^{2}-4ac &amp;gt;0\) then the graph is a hyperbola (or two intersecting lines).&lt;/p&gt;
&lt;p&gt; Consider the following equation: \(\displaystyle x^{2}+Bxy=F\).&lt;/p&gt;
&lt;br&gt;Identify the corresponding values of \(\displaystyle a=\)[[input:ans1]] \(,\displaystyle b=\)[[input:ans2]] \(,c=\) [[input:ans3a]],  \(f=\)[[input:ans3]] [[feedback:prt1]]
&lt;p&gt; [[validation:ans1]] [[validation:ans2]] [[validation:ans3]] [[validation:ans3a]]&lt;/p&gt;
&lt;p&gt; So, \(\displaystyle b^{2}−4ac=\)[[input:ans4]] \(,B \neq 0\) [[validation:ans4]] [[feedback:prt4]]&lt;/p&gt;
&lt;p&gt;Therefore, the graph of \(\displaystyle x^{2}+Bxy=F,\) is a hyperbola or two intersecting lines, if \(B \neq 0\). &lt;/p&gt;
&lt;p&gt; Case 1: If \(\displaystyle F \neq 0,\)  we have;&lt;/p&gt;
&lt;br&gt;\(\displaystyle y=\) [[input:ans5]][[validation:ans5]], thus the graph is a [[input:ans6]][[validation:ans6]][[feedback:prt5]]
&lt;p&gt; Case 2: If \(\displaystyle F = 0\)  we have;&lt;/p&gt;
&lt;p&gt;\(\displaystyle x=\)[[input:ans7]][[validation:ans7]] and \(y=\)[[input:ans8]][[validation:ans8]], thus the graph represents [[input:ans10]][[validation:ans10]]  [[feedback:prt6]]&lt;/p&gt;

&lt;p&gt;Application:&lt;/p&gt;
&lt;p&gt;What is the nature of the graph corresponding to the equation \({@eqn_1@}\)&lt;/p&gt;
&lt;p&gt;[[input:ans9]] [[validation:ans9]][[feedback:prt7]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;p&gt;&lt;h3&gt;Worked solution&lt;/h3&gt;&lt;/p&gt;
&lt;p&gt;A conic in the two-dimensional plane is the solution set of an equation of the form&lt;/p&gt;
&lt;p&gt;\(\displaystyle ax^{2}+bxy+cy^{2}+dx+ey+f=0\)&lt;/p&gt;
&lt;p&gt;for real parameters \(\displaystyle a,b,c,d,e\) and \(f,\) where at least one of the \(a,b\) and \(c\) is not zero.&lt;/p&gt;
&lt;p&gt;If \(\displaystyle b^{2}-4ac &gt;0\) then the graph is a hyperbola (or two intersecting lines).&lt;/p&gt;
&lt;p&gt; Consider the following equation: \(\displaystyle x^{2}+Bxy=F\).&lt;/p&gt;

&lt;p&gt;\(\displaystyle a={@ta_1@},b={@ta_2@},c=0,f={@ta_3@}\)&lt;/p&gt;
&lt;p&gt;\(\displaystyle b^{2}−4ac=B^{2}−4(1)(0)={@ta_4@} &gt;0,B \neq 0\).&lt;/p&gt;
&lt;p&gt;Therefore, the graph of \(\displaystyle x^{2}+Bxy=F,\) is a hyperbola or two intersecting lines, if \(B \neq 0\). &lt;/p&gt;
&lt;p&gt; Case 1: If \(\displaystyle F \neq 0,\)  we have;&lt;/p&gt;
&lt;p&gt;\(\displaystyle y={@ta_5@}\), thus the graph is a {@mcq_correct(ta_6)[1]@}.&lt;/p&gt;
&lt;p&gt; Case 2: If \(\displaystyle F = 0, x^{2}+Bxy=0,\)  we have; &lt;/p&gt;
&lt;p&gt;\(\displaystyle x={@ta_7@}\) and \(y={@ta_8@}\), thus the graph represents {@mcq_correct(ta_10)[1]@}.&lt;/p&gt;
&lt;p&gt;As a result, the graph is a hyperbola when \(F\neq 0\) and intersecting lines when \(F=0\).&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text>ta_1:1;
ta_2:B;
ta_3:-F;
ta_3a:0;
ta_4:B^2;
ta_5:(-x/B)+((F/B)/(x));
ta_7:0;
ta_8:(-1/B)*x;
ta_6: random_permutation([["hyperbola", true],["ellipse", false],["circle", false],["two intersecting lines", false]]);
var_B:rand(5)+1;
var_F:rand_with_step(-5,5,1);
eqn_1:x^2+var_B*x*y=var_F;
/*Whether it is hyperbola or not*/
ta_hyp: random_permutation([["hyperbola", true],["ellipse", false],["circle", false],["two intersecting lines", false]]);
ta_line: random_permutation([["hyperbola", false],["ellipse", false],["circle", false],["two intersecting lines", true]]);
ta_9: if var_F=0 then ta_line else ta_hyp;
ta_10: random_permutation([["hyperbola", false],["ellipse", false],["circle", false],["two intersecting lines", true]]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;\(B &gt; 0\) and \(F \neq 0 -\) hyperbola&lt;/p&gt;
&lt;p&gt;\(B &gt; 0\) and \(F = 0-\) two intersecting lines&lt;/p&gt;
&lt;p&gt;\({@eqn_1@}\)&lt;/p&gt;
&lt;p&gt;{@mcq_correct(ta_9)@}&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>dropdown</type>
      <tans>ta_10</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3a</name>
      <type>algebraic</type>
      <tans>ta_3a</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta_4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta_5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>dropdown</type>
      <tans>ta_6</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta_7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>ta_8</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>dropdown</type>
      <tans>ta_9</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans:{ans1,ans2,ans3,ans3a};
ta:{ta_1,ta_2,ta_3,ta_3a};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sans</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta_4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks ans5 and ans6 are true*/
hyper_cor: is(ans5=ta_5) and ans6="hyperbola";
/*Node 2 and 3: Checks ans5 is true and ans6 is false for partial marks*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>hyper_cor</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta_5</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>mcq_correct(ta_6)[1]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta_7</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>ta_8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.4</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt6-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>mcq_correct(ta_10)[1]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>mcq_correct(ta_9)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>684327895</deployedseed>
    <deployedseed>1175067647</deployedseed>
    <deployedseed>421362976</deployedseed>
    <deployedseed>1636100566</deployedseed>
    <deployedseed>1645302687</deployedseed>
    <deployedseed>1459253565</deployedseed>
    <deployedseed>703846294</deployedseed>
    <deployedseed>598651064</deployedseed>
    <deployedseed>884325751</deployedseed>
    <deployedseed>936506153</deployedseed>
    <deployedseed>651624280</deployedseed>
    <deployedseed>737726277</deployedseed>
    <deployedseed>252031347</deployedseed>
    <deployedseed>791653315</deployedseed>
    <deployedseed>1334805019</deployedseed>
    <deployedseed>373687508</deployedseed>
    <deployedseed>409267923</deployedseed>
    <deployedseed>2009368322</deployedseed>
    <deployedseed>189311176</deployedseed>
    <deployedseed>482148060</deployedseed>
    <deployedseed>105497164</deployedseed>
    <deployedseed>1457899673</deployedseed>
    <deployedseed>148782103</deployedseed>
    <deployedseed>1494809114</deployedseed>
    <deployedseed>652387835</deployedseed>
    <deployedseed>1070160707</deployedseed>
    <deployedseed>892440356</deployedseed>
    <deployedseed>1889075436</deployedseed>
    <deployedseed>196062505</deployedseed>
    <deployedseed>1055501452</deployedseed>
    <deployedseed>226376255</deployedseed>
    <deployedseed>1468755359</deployedseed>
    <deployedseed>206694862</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta_3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>first(mcq_correct(ta_4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
