<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Functions-(NS-and-SS)/Polynomials,-zeros-of-polynomials,-rational-functions-and-their-/Identify-the-graph-that-corresponds-to-the-given-equation-.xml">
  <question type="stack">
    <name>
      <text>Identify the graph that corresponds to the given equation.</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;(a) Select the graph with the equation \(y = {@part_a_eq@}\)&lt;/p&gt;

&lt;table&gt;
  &lt;tr style="height:225px"&gt;
    &lt;td style="width:60px; font-size:15px"&gt;
    \(\text{(i)}\)
    &lt;/td&gt;
    &lt;td style="width:250px"&gt;
    [[jsxgraph height="200px" width="200px"]]

    JXG.Options.text.useMathJax=true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox:[-5,5,5,-5],
      showCopyright:false,
      showNavigation:false,
      keepaspectratio:true
    });

    var axisStyle = {
      fixed:true,
      strokeColor:"black",
      highlightStrokeColor:"black",
      ticks:{visible:false}
    };

    board.create("axis", [[0,0],[0,1]], axisStyle);
    board.create("axis", [[0,0],[1,0]], axisStyle);

    var f1 = board.jc.snippet('{#disp_order[1][1]#}', true, 'x', true);

    board.create("functiongraph", [f1,-5, 5], {
      fixed:true,
      highlight:false,
      strokeColor:"black",
      strokeWidth:2
    });

    [[/jsxgraph]]
    &lt;/td&gt;
    &lt;td style="width:60px; font-size:15px"&gt;
    \(\text{(ii)}\)
    &lt;/td&gt;
    &lt;td style="width:250px"&gt;
    [[jsxgraph height="200px" width="200px"]]

    JXG.Options.text.useMathJax=true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox:[-5,5,5,-5],
      showCopyright:false,
      showNavigation:false,
      keepaspectratio:true
    });

    var axisStyle = {
      fixed:true,
      strokeColor:"black",
      highlightStrokeColor:"black",
      ticks:{visible:false}
    };

    board.create("axis", [[0,0],[0,1]], axisStyle);
    board.create("axis", [[0,0],[1,0]], axisStyle);

    var f2 = board.jc.snippet('{#disp_order[2][1]#}', true, 'x', true);

    board.create("functiongraph", [f2,-5, 5], {
      fixed:true,
      highlight:false,
      strokeColor:"black",
      strokeWidth:2
    });

    [[/jsxgraph]]
    &lt;/td&gt;
    &lt;/tr&gt;
    &lt;tr style="height:225px"&gt;
    &lt;td style="width:60px; font-size:15px"&gt;
    \(\text{(iii)}\)
    &lt;/td&gt;
    &lt;td style="width:250px"&gt;
    [[jsxgraph height="200px" width="200px"]]

    JXG.Options.text.useMathJax=true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox:[-5,5,5,-5],
      showCopyright:false,
      showNavigation:false,
      keepaspectratio:true
    });

    var axisStyle = {
      fixed:true,
      strokeColor:"black",
      highlightStrokeColor:"black",
      ticks:{visible:false}
    };

    board.create("axis", [[0,0],[0,1]], axisStyle);
    board.create("axis", [[0,0],[1,0]], axisStyle);

    var f3 = board.jc.snippet('{#disp_order[3][1]#}', true, 'x', true);

    board.create("functiongraph", [f3,-5, 5], {
      fixed:true,
      highlight:false,
      strokeColor:"black",
      strokeWidth:2
    });

    [[/jsxgraph]]
    &lt;/td&gt;
    &lt;td style="width:60px"&gt;
    \(\text{(iv)}\)
    &lt;/td&gt;
    &lt;td style="width:220px"&gt;
    [[jsxgraph height="200px" width="200px"]]

    JXG.Options.text.useMathJax=true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox:[-5,5,5,-5],
      showCopyright:false,
      showNavigation:false,
      keepaspectratio:true
    });

    var axisStyle = {
      fixed:true,
      strokeColor:"black",
      highlightStrokeColor:"black",
      ticks:{visible:false}
    };

    board.create("axis", [[0,0],[0,1]], axisStyle);
    board.create("axis", [[0,0],[1,0]], axisStyle);

    var f4 = board.jc.snippet('{#disp_order[4][1]#}', true, 'x', true);

    board.create("functiongraph", [f4,-5, 5], {
      fixed:true,
      highlight:false,
      strokeColor:"black",
      strokeWidth:2
    });

    [[/jsxgraph]]
    &lt;/td&gt;
    &lt;/tr&gt;
    &lt;tr style="height:225px"&gt;
    &lt;td style="width:60px"&gt;
    \(\text{(v)}\)
    &lt;/td&gt;
    &lt;td style="width:220px"&gt;
    [[jsxgraph height="200px" width="200px"]]

    JXG.Options.text.useMathJax=true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox:[-5,5,5,-5],
      showCopyright:false,
      showNavigation:false,
      keepaspectratio:true
    });

    var axisStyle = {
      fixed:true,
      strokeColor:"black",
      highlightStrokeColor:"black",
      ticks:{visible:false}
    };

    board.create("axis", [[0,0],[0,1]], axisStyle);
    board.create("axis", [[0,0],[1,0]], axisStyle);

    var f5 = board.jc.snippet('{#disp_order[5][1]#}', true, 'x', true);

    board.create("functiongraph", [f5,-5, 5], {
      fixed:true,
      highlight:false,
      strokeColor:"black",
      strokeWidth:2
    });

    [[/jsxgraph]]
    &lt;/td&gt;
    &lt;td style="width:60px"&gt;
    \(\text{(vi)}\)
    &lt;/td&gt;
    &lt;td style="width:220px"&gt;
    [[jsxgraph height="200px" width="200px"]]

    JXG.Options.text.useMathJax=true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox:[-5,5,5,-5],
      showCopyright:false,
      showNavigation:false,
      keepaspectratio:true
    });

    var axisStyle = {
      fixed:true,
      strokeColor:"black",
      highlightStrokeColor:"black",
      ticks:{visible:false}
    };

    board.create("axis", [[0,0],[0,1]], axisStyle);
    board.create("axis", [[0,0],[1,0]], axisStyle);

    var f6 = board.jc.snippet('{#disp_order[6][1]#}', true, 'x', true);

    board.create("functiongraph", [f6,-5, 5], {
      fixed:true,
      highlight:false,
      strokeColor:"black",
      strokeWidth:2
    });

    [[/jsxgraph]]
    &lt;/td&gt;
  &lt;/tr&gt;
&lt;/table&gt;

&lt;p&gt;[[input:ans1]] [[validation:ans1]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(b) In the graph below, the equation of the dashed red parabola is {@part_b_eq1@}.
&lt;br&gt; The red parabola has been translated to give the black parabola.
&lt;br&gt;Using this information, select the correct equation below for the black parabola.&lt;/p&gt;

&lt;br&gt;

&lt;p style="margin-left: 30px;"&gt;
[[jsxgraph height="400px" width="400px"]]

    JXG.Options.text.useMathJax=true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox:[
        {#b2 + x_trans/2#} - 6, 
        {#c2 + y_trans/2#} + 6, 
        {#b2 + x_trans/2#} + 6, 
        {#c2 + y_trans/2#} - 6
      ],
      showCopyright:false,
      showNavigation:false,
      keepaspectratio:true,
      grid:true
    });

    var axisStyle = {
      fixed:true,
      strokeColor:"black",
      highlightStrokeColor:"black",
      ticks:{visible:false}
    };

    board.create("axis", [[0,0],[0,1]], axisStyle);
    board.create("axis", [[0,0],[1,0]], axisStyle);

    var funcStyle = {
      fixed:true,
      highlight:false,
      strokeColor:"red",
      strokeWidth:2,
      dash:2
    };

    var f7 = board.jc.snippet('{#rhs(part_b_eq1)#}', true, 'x', true);
    board.create("functiongraph", [f7,-10,10], funcStyle);
    
    funcStyle.strokeColor = "black";
    funcStyle.dash = 1;
    var f8 = board.jc.snippet('{#rhs(part_b_eq2)#}', true, 'x', true);
    board.create("functiongraph", [f8,-10,10], funcStyle);

[[/jsxgraph]]
&lt;/p&gt;

&lt;p&gt;[[input:ans2]] [[validation:ans2]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt2]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked Solution&lt;/h3&gt;

&lt;p&gt;(a) For the first part of the question, we identify three key characteristics of the polynomial given.&lt;/p&gt;

&lt;p style="margin-left: 30px"&gt;1. The polynomial is a &lt;b&gt;quartic&lt;/b&gt;, of degree 4, so must have 3 turning points.
&lt;br&gt;This immediately excludes the two cubic equations as solutions.&lt;/p&gt;

&lt;p style="margin-left: 30px"&gt;2. The polynomial has a &lt;b&gt;{@if c1&lt;0 then "negative" else "positive"@} \(y\)-intercept&lt;/b&gt;, which we verify by calculating the value of \(y\) at \(x=0\).
&lt;br&gt;This means that the curve should cross the \(y\)-axis in the {@if c1&lt;0 then "lower" else "upper"@} half-plane.&lt;/p&gt;

&lt;p style="margin-left: 30px"&gt;3. The polynomial has a &lt;b&gt;{@if a1&lt;0 then "negative" else "positive"@} limit&lt;/b&gt; as \(x \rightarrow \pm \inf\)., which we verify by considering the sign of the coefficient of the highest power term, \(x^4\).
&lt;br&gt;This means that the tails of the curve should therefore {@if a1&lt;0 then "decrease" else "increase"@} away from the origin.&lt;/p&gt;

&lt;p&gt;From these three observations, only one of the curves has the required characteristics to represent the equation \(y = {@part_a_eq@}\)
&lt;br&gt;The correct answer is therefore {@
    if disp_order[1][2] then m_ans1[1][3]
    else if disp_order[2][2] then m_ans1[2][3]
    else if disp_order[3][2] then m_ans1[3][3]
    else if disp_order[4][2] then m_ans1[4][3]
    else if disp_order[5][2] then m_ans1[5][3]
    else if disp_order[6][2] then m_ans1[6][3]
@}.&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(b) In the second part of the question, we need to determine the translation which has been applied to transform from the red to the black parabola.
&lt;br&gt;The simplest way to do this is to consider how the turning point of the parabola has been shifted.&lt;/p&gt;

&lt;p&gt;Counting the squares gives that the parabola has been translated by,&lt;/p&gt;
&lt;p style="margin-left: 30px" &gt;\({@abs(x_trans)@}\) squares {@if x_trans&gt;0 then "rightward" else "leftward"@}&lt;/p&gt;
&lt;p style="margin-left: 30px" &gt;\({@abs(y_trans)@}\) squares {@if y_trans&gt;0 then "upward" else "downward"@}&lt;/p&gt;
&lt;p&gt;Or, by a vector of \( ({@x_trans@}, {@y_trans@}) \).&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;We may then write the formula for the black parabola, using the fact that the result of applying a translation \( (x_0, y_0) \) to a curve with equation \(y = f(x)\) is given by,&lt;/p&gt;
&lt;p style="margin-left: 30px"&gt;\(y - y_0= f(x - x_0)\)&lt;/p&gt;

&lt;p hidden&gt;{@simp:false;@}&lt;/p&gt;
&lt;p&gt;The equation for the black parabola is therefore given by,&lt;/p&gt;
&lt;p style="margin-left: 30px"&gt;\({@y - y_trans@} = {@a2*(x - x_trans)^2 - ev(2*a2*b2,simp)*(x - x_trans) + ev(a2*(b2^2) + c2,simp)@}\)&lt;/p&gt;
&lt;p style="margin-left: 30px"&gt;\(y = {@a2*(x - x_trans)^2 - ev(2*a2*b2,simp)*(x - x_trans) + ev(a2*(b2^2) + c2 + y_trans, simp)@}\)&lt;/p&gt;
&lt;p style="margin-left: 30px"&gt;\(y = {@rhs(part_b_eq2)@}\)&lt;/p&gt;
&lt;p hidden&gt;{@simp:true;@}&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>5</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Part (a) - randomizing equation type and display order */

eq_type: rand(4)+1;
/* 1: +ve limits, +ve y-intercept
   2: +ve limits, -ve y-intercept
   3: -ve limits, +ve y-intercept
   4: -ve limits, -ve y-intercept */

disp_order: random_permutation([
  [(1/10)*x^4 - x^2 + 1, is(eq_type=1)],
  [(1/10)*x^4 - x^2 - 1, is(eq_type=2)],
  [-(1/10)*x^4 + x^2 + 1, is(eq_type=3)],
  [-(1/10)*x^4 + x^2 - 1, is(eq_type=4)],
  [(1/4)*x^3 - 2*x + 1, false],
  [-(1/4)*x^3 + 2*x - 1, false]
]);

m_ans1: [
  [disp_order[1][1], disp_order[1][2], "\\(\\text{(i)}\\)"],
  [disp_order[2][1], disp_order[2][2], "\\(\\text{(ii)}\\)"],
  [disp_order[3][1], disp_order[3][2], "\\(\\text{(iii)}\\)"],
  [disp_order[4][1], disp_order[4][2], "\\(\\text{(iv)}\\)"],
  [disp_order[5][1], disp_order[5][2], "\\(\\text{(v)}\\)"],
  [disp_order[6][1], disp_order[6][2], "\\(\\text{(vi)}\\)"]
];

t_ans1: mcq_correct(m_ans1)[1];

/* Part (a) - Generating equation */

a1: (if elementp(eq_type,{1,2}) then 1 else -1)*(rand(5)+1);
c1: (if elementp(eq_type,{1,3}) then 1 else -1)*(rand(5)+1);

b1: if eq_type = 1 then
    rand_with_step(ceiling(-2*sqrt(4*a1*c1)), floor(-sqrt(4*a1*c1)), 1)
  else if eq_type = 2 then
    rand_with_step(-3*a1, -a1, 1)
  else if eq_type = 3 then
    rand_with_step(-a1, -3*a1, 1)
  else if eq_type = 4 then
    rand_with_step(ceiling(sqrt(4*a1*c1)), floor(2*sqrt(4*a1*c1)), 1);

part_a_eq: a1*x^4 + b1*x^2 + c1;

/* Part (b) - randomizing start parabola */

a2: rand_with_prohib(-1,1,[0]);
b2: rand_with_prohib(-4,4,[0]);
c2: rand_with_prohib(-4,4,[0]);

part_b_eq1: y = expand(a2*(x - b2)^2 + c2);

/* Part (b) - randomizing translation (towards centre) */

trans_vec: random_permutation([(rand(4)+3), (rand(3)+1)])
x_trans: (if b2 &lt; 0 then 1 else -1)*trans_vec[1];
y_trans: (if c2 &lt; 0 then 1 else -1)*trans_vec[2];

part_b_eq2: y = expand(a2*(x - b2 - x_trans)^2 + c2 + y_trans);

/* Part (b) - randomizing order of solutions */

m_ans2: random_permutation([
  [part_b_eq2, true],
  [y = expand(a2*(x - b2 + x_trans)^2 + c2 + y_trans), false],
  [y = expand(a2*(x - b2 - x_trans)^2 + c2 - y_trans), false],
  [y = expand(a2*(x - b2 + x_trans)^2 + c2 - y_trans), false]
])</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;(a) {@part_a_eq@}&lt;/p&gt;
&lt;p&gt;(b) Start equation: {@part_b_eq1@}
&lt;br&gt;End equation: {@part_b_eq2@}&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>m_ans1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>latex</options>
    </input>
    <input>
      <name>ans2</name>
      <type>radio</type>
      <tans>m_ans2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1 - Check correct answer */

/* Node 2 - Check if answer is cubic equation */
node2_test: if ans1=(1/4)*x^3-2*x+1 or ans1=-(1/4)*x^3+2*x-1 then 1 else 0;

/* Node 3 - Check if answer has wrong y-intercept */

incorr_intercept: if c1 &gt; 0 then 
    {(1/10)*x^4 - x^2 - 1, -(1/10)*x^4 + x^2 - 1}
  else if c1 &lt; 0 then
    {(1/10)*x^4 - x^2 + 1, -(1/10)*x^4 + x^2 + 1};

node3_test: if elementp(ans1, incorr_intercept) then 1 else 0;

/* Node 4 - Check if answer has wrong limit behaviour */

incorr_limit: if a1 &lt; 0 then 
    {(1/10)*x^4 - x^2 + 1, (1/10)*x^4 - x^2 - 1}
  else if a1 &gt; 0 then
    {-(1/10)*x^4 + x^2 + 1, -(1/10)*x^4 + x^2 - 1};

node4_test: if elementp(ans1, incorr_limit) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>t_ans1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node2_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;The answer which you have given has the shape of a &lt;b&gt;cubic&lt;/b&gt; polynomial, of degree 3.
&lt;br&gt;This can be identified by the fact that the curve has 2 turning points.&lt;/p&gt;

&lt;p&gt;However, the equation given represents a &lt;b&gt;quartic&lt;/b&gt; polynomial, of degree 4.
&lt;br&gt;The curve should therefore have 3 turning points.&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;The answer which you have given has a &lt;b&gt;{@if c1&lt;0 then "positive" else "negative"@} \(y\)-intercept&lt;/b&gt;.
&lt;br&gt;This can be identified by the fact that the curve crosses the \(y\)-axis in the {@if c1&lt;0 then "upper" else "lower"@} half-plane.&lt;/p&gt;

&lt;p&gt;However, the equation given represents a curve with a &lt;b&gt;{@if c1&lt;0 then "negative" else "positive"@} \(y\)-intercept&lt;/b&gt;.
&lt;br&gt;This can be seen by checking the \(y\)-value of the curve at \(x=0\). In this case, we have \(y = {@c1@}\).
&lt;br&gt;The curve should therefore cross the \(y\)-axis in the {@if c1&lt;0 then "lower" else "upper"@} half-plane.&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node4_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text>{@if node3_test=1 then "&lt;br&gt;" else ""@}

&lt;p&gt;The answer which you have given has a &lt;b&gt;{@if a1&lt;0 then "positive" else "negative"@} limit&lt;/b&gt; as \(x \rightarrow \pm \inf\).
&lt;br&gt;This can be identified by the fact that the tails of the curve {@if a1&lt;0 then "increase" else "decrease"@} away from the origin.&lt;/p&gt;

&lt;p&gt;However, the equation given represents a curve with a &lt;b&gt;{@if a1&lt;0 then "negative" else "positive"@} limit&lt;/b&gt; as \(x \rightarrow \pm \inf\).
&lt;br&gt;This can be seen by checking the sign of the coefficient of the highest power term, \(x^4\). In this case, we have \({@if a1&lt;0 then a1&lt;0 else a1&gt;0@}\).
&lt;br&gt;The tails of the curve should therefore {@if a1&lt;0 then "decrease" else "increase"@} away from the origin.&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1 - Check correct answer */

/* Node 2 - Check whether transformed in wrong x-direction */

incorr_x_trans: {
    y = expand(a2*(x - b2 + x_trans)^2 + c2 + y_trans), 
    y = expand(a2*(x - b2 + x_trans)^2 + c2 - y_trans)
  };

node2_test: if elementp(ans2, incorr_x_trans) then 1 else 0;

/* Node 3 - Check whether transformed in wrong y-direction */

incorr_y_trans: {
    y = expand(a2*(x - b2 - x_trans)^2 + c2 - y_trans),
    y = expand(a2*(x - b2 + x_trans)^2 + c2 - y_trans)
  };

node3_test: if elementp(ans2, incorr_y_trans) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>part_b_eq2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node2_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;The equation which you have selected corresponds to a translation in the &lt;b&gt;{@if x_trans&lt;0 then "positive" else "negative"@}&lt;/b&gt; \(x\)-direction.
&lt;br&gt;However, in the graph, the red parabola has been translated {@if x_trans&lt;0 then "leftwards" else "rightwards"@}, in the &lt;b&gt;{@if x_trans&lt;0 then "negative" else "positive"@}&lt;/b&gt; \(x\)-direction.
&lt;br&gt;Remember that the translation of a graph \(y = f(x)\) in the \(x\)-direction, by a vector \( (x_0, 0) \), is given by,&lt;/p&gt; 

&lt;p style="margin-left: 30px"&gt;\(y = f(x - x_0)\)&lt;/p&gt;

&lt;p hidden&gt;{@simp: false;@}&lt;/p&gt;
&lt;p&gt;In this case, the graph has been translated by \( ({@x_trans@}, {@y_trans@}) \).
&lt;br&gt;So, the result of the translation in the \(x\)-direction &lt;b&gt;only&lt;/b&gt; is \(y = {@a2*(x - x_trans)^2 - ev(2*a2*b2,simp)*(x - x_trans) + ev(a2*(b2^2) + c2,simp)@}\).&lt;/p&gt;
&lt;p hidden&gt;{@simp: true;@}&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text>{@if node2_test=1 then "&lt;br&gt;" else ""@}

&lt;p&gt;The equation which you have selected corresponds to a translation in the &lt;b&gt;{@if y_trans&lt;0 then "positive" else "negative"@}&lt;/b&gt; \(y\)-direction.
&lt;br&gt;However, in the graph, the red parabola has been translated {@if y_trans&lt;0 then "downwards" else "upwards"@}, in the &lt;b&gt;{@if y_trans&lt;0 then "negative" else "positive"@}&lt;/b&gt; \(x\)-direction.
&lt;br&gt;Remember that the translation of a graph \(y = f(x)\) in the \(y\)-direction, by a vector \( (0, y_0) \), is given by,&lt;/p&gt;

&lt;p style="margin-left: 30px"&gt;\(y - y_0 = f(x)\)&lt;/p&gt;

&lt;p hidden&gt;{@simp: false;@}&lt;/p&gt;
&lt;p&gt;In this case, the graph has been translated by \( ({@x_trans@}, {@y_trans@}) \).
&lt;br&gt;So, the result of the translation in the \(y\)-direction &lt;b&gt;only&lt;/b&gt; is \({@y - y_trans = ev(expand(a2*(x - b2)^2 + c2), simp)@}\).&lt;/p&gt;
&lt;p hidden&gt;{@simp: true;@}&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1962631295</deployedseed>
    <deployedseed>1109218150</deployedseed>
    <deployedseed>1031077491</deployedseed>
    <deployedseed>65043273</deployedseed>
    <deployedseed>930238107</deployedseed>
    <deployedseed>635492929</deployedseed>
    <deployedseed>2056365451</deployedseed>
    <deployedseed>900310536</deployedseed>
    <deployedseed>2053208296</deployedseed>
    <deployedseed>2115346038</deployedseed>
    <deployedseed>1209243096</deployedseed>
    <deployedseed>1184155612</deployedseed>
    <deployedseed>426714428</deployedseed>
    <deployedseed>1166108936</deployedseed>
    <deployedseed>1637936160</deployedseed>
    <deployedseed>2090931000</deployedseed>
    <deployedseed>2129190474</deployedseed>
    <deployedseed>1744295380</deployedseed>
    <deployedseed>1022855344</deployedseed>
    <deployedseed>919923837</deployedseed>
    <deployedseed>415058775</deployedseed>
    <deployedseed>2052735585</deployedseed>
    <deployedseed>754694844</deployedseed>
    <deployedseed>1643269001</deployedseed>
    <deployedseed>374032286</deployedseed>
    <deployedseed>2037371921</deployedseed>
    <deployedseed>378003151</deployedseed>
    <deployedseed>1180669832</deployedseed>
    <deployedseed>925862186</deployedseed>
    <deployedseed>646586480</deployedseed>
    <deployedseed>598854749</deployedseed>
    <deployedseed>1837194407</deployedseed>
    <deployedseed>2106063187</deployedseed>
    <deployedseed>628682423</deployedseed>
    <deployedseed>744074662</deployedseed>
    <deployedseed>2142276837</deployedseed>
    <deployedseed>696460982</deployedseed>
    <deployedseed>773081634</deployedseed>
    <deployedseed>103583069</deployedseed>
    <deployedseed>2114881545</deployedseed>
    <deployedseed>799391812</deployedseed>
    <deployedseed>1173200701</deployedseed>
    <deployedseed>358977678</deployedseed>
    <deployedseed>809956270</deployedseed>
    <deployedseed>1915397801</deployedseed>
    <deployedseed>785287974</deployedseed>
    <deployedseed>673968026</deployedseed>
    <deployedseed>314392599</deployedseed>
    <deployedseed>1557050162</deployedseed>
    <deployedseed>295587446</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(m_ans1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(m_ans2))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
