<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Propositional-Logic-and-Set-Theory-(NS-and-SS)/Definition-and-examples-of-propositions/Propositional-definitions.xml">
  <question type="stack">
    <name>
      <text>Propositional definitions</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text>&lt;p&gt;Complete the following definitions relating to propositional logic.&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(a) A “proposition” is a statement that is [[input:ans1]]&lt;/p&gt;
&lt;p&gt;[[validation:ans1]]&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(b) The “implication”, \(p \Rightarrow q\), is the statement that,&lt;/p&gt;
&lt;p&gt;if \(p\) is true [[input:ans2]] \(q\) is [[input:ans3]].&lt;/p&gt;
&lt;p&gt;[[validation:ans2]] [[validation:ans3]]&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(c) The “bi-implication” or “equivalence”, \(p \iff q\), is the statement that,&lt;/p&gt;
&lt;p&gt;\(p\) is true [[input:ans4]] \(q\) is also [[input:ans5]].&lt;/p&gt;
&lt;p&gt;[[validation:ans4]] [[validation:ans5]]&lt;/p&gt;

&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked solution&lt;/h3&gt;

&lt;p&gt;The correct completed statements are the following:&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(a) A “proposition” is a statement that is &lt;b&gt;either true or false but not both&lt;/b&gt;.
&lt;br&gt;For example, “a cube has 9 faces” or “\(x^2 - 3x + 2 = 0\)".&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(b) The “implication”, \(p \Rightarrow q\), means &lt;b&gt;if \(p\) is true then \(q\) is true&lt;/b&gt;.
&lt;br&gt;This does not mean that the truth of \(q\) depends on \(p\); it is simply the logical relationship “if \(p\), then \(q\)”.
&lt;br&gt;For example, “if \(x=2\), then \(x^2 - 3x + 2 = 0\)”.&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(c) The “bi-implication” or “equivalence”, \(p \iff q\), means that &lt;b&gt;\(p\) is true if and only if \(q\) is true&lt;/b&gt;.
&lt;br&gt;For example, “\(x^2 - 3x + 2 = 0\) if and only if \(x=2\) or \(x=1\)”.&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Randomization of incorrect options */

/* (a) */
incorrect_ops_a: rand_selection([
  ["true", false],
  ["false", false],
  ["indeterminable", false],
  ["both true and false", false],
  ["neither true nor false", false]
], 2);
m_ans1: random_permutation(append(
  [["either true or false but not both", true]],
  [["either true or false", false]],
  incorrect_ops_a
));
t_ans1: "either true or false but not both";

/* (b) – Fix: implication is “if p then q” */
m_ans2: random_permutation([
  ["if", true],
  ["only if", false],
  ["if and only if", false],
  ["depending on whether", false],
  ["exactly when", false]
]);
t_ans2: "if";

m_ans3: random_permutation([
  ["true", true],
  ["false", false],
  ["indeterminable", false]
]);
t_ans3: "true";

/* (c) – Fix: bi-implication is “if and only if p and q are true” */
m_ans4: random_permutation([
  ["if and only if", true],
  ["if", false],
  ["only if", false],
  ["depending on whether", false]
]);
t_ans4: "if and only if";

m_ans5: random_permutation([
  ["true", true],
  ["false", false],
  ["indeterminable", false]
]);
t_ans5: "true";
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;Part (a) options: {@m_ans1@}&lt;/p&gt;
&lt;p&gt;Part (b)(i) options: {@m_ans2@}&lt;/p&gt;
&lt;p&gt;Part (b)(ii) options: {@m_ans3@}&lt;/p&gt;
&lt;p&gt;Part (c)(i) options: {@m_ans4@}&lt;/p&gt;
&lt;p&gt;Part (c)(ii) options: {@m_ans5@}&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>m_ans1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>m_ans2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>m_ans3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>m_ans4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>dropdown</type>
      <tans>m_ans5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1: Check all answers correct */
node1_test: if [ans1, ans2, ans3, ans4, ans5] = [t_ans1, t_ans2, t_ans3, t_ans4, t_ans5] then 1 else 0;

/* Node 2: Award one mark if two of the three statements are correct */
correct_part_a: is(ans1=t_ans1);
correct_part_b: is([ans2, ans3]=[t_ans2, t_ans3]);
correct_part_c: is([ans4, ans5]=[t_ans4, t_ans5]);
node2_test: if (correct_part_a and correct_part_b) or (correct_part_a and correct_part_c) or (correct_part_b and correct_part_c) then 1 else 0;

/* Node 3: Feedback for part (a) */
node3_test: if correct_part_a then 1 else 0;

/* Node 4: Feedback for part (b) */
node4_test: if correct_part_b then 1 else 0;

/* Node 5: Feedback for part (c) */
node5_test: if correct_part_c then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node1_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node2_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Your definition of a proposition is not correct.
&lt;br&gt; A &lt;b&gt;proposition&lt;/b&gt; is a statement that is &lt;b&gt;either true or false but not both&lt;/b&gt;.&lt;/p&gt;</text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node4_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Your definition of an implication is not correct.
&lt;br&gt; An &lt;b&gt;implication&lt;/b&gt; is a statement that a proposition is true &lt;b&gt;if&lt;/b&gt; or &lt;b&gt;only if&lt;/b&gt; another proposition is &lt;b&gt;true&lt;/b&gt;.&lt;/p&gt;</text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node5_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Your definition of a bi-implication is not correct.
&lt;br&gt; A &lt;b&gt;bi-implication&lt;/b&gt; is a statement that a proposition is true &lt;b&gt;if and only if&lt;/b&gt; another proposition is &lt;b&gt;true&lt;/b&gt;.&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1350179340</deployedseed>
    <deployedseed>1606511609</deployedseed>
    <deployedseed>28621883</deployedseed>
    <deployedseed>2112990057</deployedseed>
    <deployedseed>706573573</deployedseed>
    <deployedseed>2099196757</deployedseed>
    <deployedseed>525350023</deployedseed>
    <deployedseed>801105833</deployedseed>
    <deployedseed>120060253</deployedseed>
    <deployedseed>1850836350</deployedseed>
    <deployedseed>781605427</deployedseed>
    <deployedseed>1469831423</deployedseed>
    <deployedseed>408561614</deployedseed>
    <deployedseed>1383276995</deployedseed>
    <deployedseed>495929505</deployedseed>
    <deployedseed>438545378</deployedseed>
    <deployedseed>1995158778</deployedseed>
    <deployedseed>1231072045</deployedseed>
    <deployedseed>1273957123</deployedseed>
    <deployedseed>1878126588</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>t_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>t_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>t_ans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>t_ans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>t_ans5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>"either true or false"</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>t_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>t_ans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>t_ans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>t_ans5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-5-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>t_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>incorrect_ops_b_i[1][1]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>t_ans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>t_ans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>t_ans5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-5-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>t_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>t_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>"false"</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>t_ans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>t_ans5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-5-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>t_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>t_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>t_ans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>incorrect_ops_c_i[1][1]</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>t_ans5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-5-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>t_ans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>t_ans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>t_ans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>t_ans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>"false"</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-5-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
