<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/The-Real-and-Complex-Number-Systems-(NS)/The-set-of-natural-numbers/Proving-the-sum-of-sequence-of-natural-numbers-using-mathematica.xml">
  <question type="stack">
    <name>
      <text>Proving the sum of sequence of natural numbers using mathematical induction</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Consider the following proposition, then fill in the blanks to complete the proof outlined below it.&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;&lt;strong&gt;Proposition:&lt;/strong&gt; Show that by mathematical induction that \( 1+2+3+\cdots+k = \frac{k(k+1)}{2}\) &lt;/p&gt;
&lt;p&gt;&lt;em&gt;Note: "LHS" means Left Hand Side and "RHS" means Right Hand Side.&lt;/em&gt;&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;&lt;strong&gt;Proof by induction:&lt;/strong&gt;&lt;/p&gt;
&lt;p&gt;Define \(P(k)\) to be the statement of the proposition, that: \(1+2+3+\cdots+k = \frac{k(k+1)}{2} \)&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;&lt;b&gt;Step 1:&lt;/b&gt; Base case.&lt;/p&gt;
&lt;p&gt;For \(k = 1\), we have&lt;/p&gt;

&lt;p style="margin-left:30px"&gt;LHS: \( 1+2+3+\cdots+k=\) [[input:ans1]] [[validation:ans1]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt1]]&lt;/ul&gt;

&lt;p style="margin-left:30px"&gt;RHS: \( \frac{k(k+1)}{2}=\) [[input:ans2]] [[validation:ans2]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt2]]&lt;/p&gt;

&lt;p&gt;Since these are equal, \( P(1) \) is [[input:ans3]] [[validation:ans3]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt3]]&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;&lt;b&gt;Step 2:&lt;/b&gt; Induction step.&lt;/p&gt;
&lt;p&gt;Assume \(P(k)\) is true, which means that \(1+2+3+\cdots+k = \frac{k(k+1)}{2} \)
&lt;br&gt;This is called the &lt;b&gt;induction hypothesis&lt;/b&gt;.&lt;/p&gt;
&lt;p&gt;We then need to show \( P(k+1) \) is true.&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;\(P(k+1): 1 + 2 + 3 + \cdots + k \ +\) [[input:ans4]] [[validation:ans4]] \(=\) [[input:ans5]] [[validation:ans5]] &amp;nbsp; (using the induction hypothesis)&lt;/p&gt;
&lt;p style="margin-left:321px"&gt;\(=\) [[input:ans5b]] [[validation:ans5b]] &amp;nbsp; (simplifying the result)&lt;/p&gt;
&lt;p&gt;[[feedback:prt4]]&lt;/p&gt;

&lt;p&gt;So, \(P(k+1)\) is [[input:ans6]] [[validation:ans6]] [[feedback:prt5]]&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;Therefore, by the Principle of Mathematical Induction, we have \(1+2+3+\cdots+k = \frac{k(k+1)}{2}\)&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;p&gt;&lt;strong&gt;Proposition:&lt;/strong&gt; Show that by mathematical induction that \( 1+2+3+\cdots+k = \frac{k(k+1)}{2}\) &lt;/p&gt;

&lt;br&gt;

&lt;p&gt;&lt;strong&gt;Proof by induction:&lt;/strong&gt;&lt;/p&gt;
&lt;p&gt;Define \(P(k)\) to be the statement of the proposition, that: \(1+2+3+\cdots+k = \frac{k(k+1)}{2} \)&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;&lt;b&gt;Step 1:&lt;/b&gt; Base case.&lt;/p&gt;
&lt;p&gt;For \(k = 1\), we have&lt;/p&gt;
&lt;p style="margin-left:30px"&gt;LHS: \( 1+2+3+\cdots+k=1\)&lt;/p&gt;
&lt;p style="margin-left:30px"&gt;RHS: \( \frac{k(k+1)}{2}=\frac{1 \cdot (1+1)}{2} = 1\)
&lt;p&gt;Since these are equal, \( P(1) \) is &lt;b&gt;true&lt;/b&gt;.&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;&lt;b&gt;Step 2:&lt;/b&gt; Induction step.&lt;/p&gt;
&lt;p&gt;Assume \(P(k)\) is true, which means that \(1+2+3+\cdots+k = \frac{k(k+1)}{2} \)
&lt;br&gt;This is called the &lt;b&gt;induction hypothesis&lt;/b&gt;.&lt;/p&gt;
&lt;p&gt;We then need to show \( P(k+1) \) is true.&lt;/p&gt;

&lt;p&gt;\(\begin{align*} 
P(k+1): 1 + 2 + 3 + \cdots + k + (k+1) &amp;= \frac{k(k+1)}{2} + (k+1) \\
&amp;= \frac{(k+1)(k+2)}{2}
\end{align*}\)&lt;/p&gt;
&lt;p&gt;Where we have used the induction hypothesis in the first equality and simplified the result in the second equality.&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;So, \(P(k+1)\) is &lt;b&gt;true&lt;/b&gt;.&lt;/p&gt;
&lt;p&gt;Therefore, by the Principle of Mathematical Induction, we have \(1+2+3+\cdots+k = \frac{k(k+1)}{2}\)&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>ta1:1;
ta2:1;
ta3:true;
sm:sum(k,k,1,k);
ta4:k+1;
ta5:k*(k+1)/2 + (k+1);
ta5b:(k+1)*(k+2)/2;
ta6:true;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;\({@ta1@}\)&lt;/p&gt;
&lt;p&gt;\({@ta2@}\)&lt;/p&gt;
&lt;p&gt;\({@ta3@}\)&lt;/p&gt;
&lt;p&gt;\({@ta4@}\)&lt;/p&gt;
&lt;p&gt;\({@ta5@}\)&lt;/p&gt;
&lt;p&gt;\({@ta6@}\)&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>boolean</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5b</name>
      <type>algebraic</type>
      <tans>ta5b</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>boolean</type>
      <tans>ta6</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Since \(n=1\), we only have the first term of the series.
&lt;br&gt;So, the series evaluates to \(1\) for the base case.&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Evaluating the expression at \(n=1\) gives, \( \frac{n(n+1)}{2}=\frac{1 \cdot (1+1)}{2} = 1\)&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Since the LHS and RHS of the equation in \(P(1)\) are equal to \(1\), we have that the statement of \(P(1)\) is &lt;b&gt;true&lt;/b&gt;.&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* node 1: checks correct LHS of P(n+1) */

/* node 2: checks correct substitution of induction hypothesis */
/* node 3: checks correct node 2 but wrong form */

/* node 4: checks correct simplification */
/* node 5: checks correct node 4 but wrong form */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/4</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;The &lt;b&gt;LHS of \(P(n+1)\)&lt;/b&gt; is the sum of the integers from \(1\) to \(n+1\).
&lt;br&gt;So, the final term in the series is \(n+1\).&lt;/p&gt;</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/4</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@if ans4=ta4 then "" else "&lt;br&gt;"@}</text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt4-3-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;The answer you have given for the &lt;b&gt;RHS of the first line&lt;/b&gt; is equal to the series given.
&lt;br&gt;However, your answer is not of the required form.
&lt;br&gt;You should have only used the induction hypothesis to replace the sum on the LHS,&lt;/p&gt;

&lt;p style="margin-left:30px"&gt;\(1 + 2 + \cdots + n = \frac{n(n+1)}{2}\)&lt;/p&gt;

&lt;p&gt;The correct answer is therefore, \(\frac{n(n+1)}{2} + (n+1)\)&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt4-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;The answer you have given for the &lt;b&gt;RHS of the first line&lt;/b&gt; is not equal to the series given.
&lt;br&gt;You should have only used the induction hypothesis to replace the sum on the LHS,&lt;/p&gt;

&lt;p style="margin-left:30px"&gt;\(1 + 2 + \cdots + n = \frac{n(n+1)}{2}\)&lt;/p&gt;

&lt;p&gt;The correct answer is therefore, \(\frac{n(n+1)}{2} + (n+1)\)&lt;/p&gt;</text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans5b</sans>
        <tans>ta5b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/2</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt4-4-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@if ans4=ta4 and ans5=ta5 then "" else "&lt;br&gt;"@}</text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5b</sans>
        <tans>ta5b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-5-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;The answer you have given for the &lt;b&gt;RHS of the second line&lt;/b&gt; is equal to the series given.
&lt;br&gt;However, your answer is not of the required form.
&lt;br&gt;You should have simplified the line above to be equal to the RHS of the statement \(P(n+1)\)&lt;/p&gt;

&lt;p style="margin-left:30px"&gt;\(\begin{align*}
\textstyle \frac{n(n+1)}{2} + (n+1) &amp;= (n+1)(\frac{n}{2} + 1) \\
\textstyle &amp;= (n+1)(\frac{n+2}{2}) \\
\textstyle &amp;= \frac{(n+1)(n+2)}{2}
\end{align*}\)&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-5-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;The answer you have given for the &lt;b&gt;RHS of the second line&lt;/b&gt; is not equal to the series given.
&lt;br&gt;However, your answer is not of the required form.
&lt;br&gt;You should have simplified the line above to be equal to the RHS of the statement \(P(n+1)\)&lt;/p&gt;

&lt;p style="margin-left:30px"&gt;\(\begin{align*}
\textstyle \frac{n(n+1)}{2} + (n+1) &amp;= (n+1)(\frac{n}{2} + 1) \\
\textstyle &amp;= (n+1)(\frac{n+2}{2}) \\
\textstyle &amp;= \frac{(n+1)(n+2)}{2}
\end{align*}\)&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;\(P(n+1)\) is the statement that \(1 + 2 + \cdots + n + (n+1) = \frac{(n+1)(n+2)}{2}\)&lt;/p&gt;
&lt;p&gt;We have shown that this is true using the induction hypothesis in the previous part.
&lt;br&gt;Therefore, \(P(n+1)\) is &lt;b&gt;true&lt;/b&gt;.&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
