<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Propositional-Logic-and-Set-Theory-(NS-and-SS)/The-concept-of-a-set/Relationship-between-elements-and-sets-notation.xml">
  <question type="stack">
    <name>
      <text>Relationship between elements and sets  notation</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text>&lt;p&gt;Answer the following questions to get a better understanding of the set theory concept.&lt;/p&gt;
 &lt;p&gt;1. If \(A\) is a set and \(x\) is an object, how do we use notation to signify whether \(x\) is an element of set \(A\) or not?&lt;/p&gt;
&lt;p&gt; a) To show \(x\) is an element of \(A\) [[input:ans1]] [[validation:ans1]]&lt;/p&gt;
 &lt;p&gt;b) To show  \(x\) is not an element of \(A\) [[input:ans2]][[validation:ans2]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;
&lt;p&gt;2. Consider the sets \(A\) and \(B\):&lt;/p&gt;
&lt;p&gt;Set \(A\) contains prime numbers less than \({@selected_number@}\) .&lt;/p&gt;
&lt;p&gt;Set \(B\) contains the numbers \(\{{@var_even_set@}\}\).&lt;/p&gt;
&lt;p&gt;&lt;p&gt;For each statement below, choose the suitable set notation that correctly describes the given situation involving the sets.&lt;/p&gt;
&lt;/p&gt;
[[input:ans3]][[validation:ans3]][[feedback:prt3]]</text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text>&lt;h3&gt;Worked solution&lt;/h3&gt; 

&lt;ol type="a"&gt; 
&lt;li&gt;
&lt;p&gt;If \(A\) is a set and \(x\) is an object, we use specific notation to show whether \(x\) belongs to \(A\) or not.&lt;/p&gt;
      &lt;ul&gt;
            &lt;li&gt;
                   &lt;p&gt;To show that \(x\) &lt;b&gt;is an element &lt;/b&gt;of \(A\), we write&lt;/p&gt;
                    &lt;p&gt;\(x \in A \) &lt;/p&gt;
                   &lt;p&gt;The symbol "\(\in \)" means "is an element of."&lt;/p&gt;
            &lt;/li&gt;
            &lt;li&gt;
                   &lt;p&gt;To show that \(x\) is &lt;b&gt;not an element&lt;/b&gt; of \(A\), we write&lt;/p&gt;
                    &lt;p&gt;\(x \notin A \) &lt;/p&gt;
                   &lt;p&gt;The symbol "\(\notin \)" means "is not an element of."&lt;/p&gt;
            &lt;/li&gt;     
   &lt;/ul&gt;
&lt;/li&gt;

&lt;li&gt;
&lt;p&gt;Consider the sets \(A\) and \(B\). Set \(A\) contains prime numbers less than {@selected_number@} . Set \(B\) contains the numbers \(\{{@var_even_set@}\}\).&lt;/p&gt;
&lt;ul&gt;
&lt;li&gt;
&lt;p&gt;{@element_to_check@} \( \in A\): This is &lt;b&gt;true&lt;/b&gt;, because {@element_to_check@} is a &lt;b&gt;prime number less than {@selected_number@}&lt;/b&gt; and therefore it is an element of \(A\).&lt;/p&gt;
&lt;/li&gt;
&lt;li&gt;
&lt;p&gt;{@element_to_check_1@} \( \notin A\): This is &lt;b&gt;true&lt;/b&gt;, because {@element_to_check_1@} is &lt;b&gt;not a prime number&lt;/b&gt; and therefore it does not belong to \(A\).&lt;/p&gt;
&lt;/li&gt;
&lt;li&gt;
&lt;p&gt;\(A \subset B\): This is &lt;b&gt;false&lt;/b&gt;, because not all elements of \(A\) (which are primes) are contained in \(B\) (which contains only even numbers).&lt;/p&gt;
&lt;/li&gt;
&lt;li&gt;
&lt;p&gt;\(B \subset A\): This is &lt;b&gt;false&lt;/b&gt;, because no even number greater than 2 is a prime number less than {@selected_number@}.&lt;/p&gt;
&lt;/li&gt;
&lt;/ul&gt;
&lt;/li&gt;
&lt;/ol&gt;</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Prime numbers */
prime_numbers: [2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47];

/* Randomly select an upper bound */
random_index: rand(length(prime_numbers)) + 1;
selected_number: prime_numbers[random_index];

/* Build set of primes less than selected_number */
prime_numbers_less_than_selected: [];
for num in prime_numbers do (
    if num &lt; selected_number then (
        prime_numbers_less_than_selected: endcons(num, prime_numbers_less_than_selected)
    )
);

/* Randomly select an element to check from prime_numbers_less_than_selected */
if length(prime_numbers_less_than_selected) &gt; 0 then (
    p_check_index: rand(length(prime_numbers_less_than_selected)) + 1,
    element_to_check: prime_numbers_less_than_selected[p_check_index]
) else (
    element_to_check: "none"
);

var_prime: prime_numbers_less_than_selected;
var_member: if member(element_to_check, prime_numbers_less_than_selected) then
    "\\(" + string(element_to_check) + " \\in A\\)"
else
    "\\(" + string(element_to_check) + " \\notin A\\)";

/* Even numbers */
random_numbers: [10, 20, 30, 40, 50];
random_index: rand(length(random_numbers)) + 1;
selected_number_1: random_numbers[random_index];

/* Build set of even numbers below selected_number_1 */
even_numbers: [];
for num: 4 step 2 while num &lt; selected_number_1 do (
    even_numbers: endcons(num, even_numbers)
);

/* Randomly select an element to check from even_numbers */
if length(even_numbers) &gt; 0 then (
    e_check_index: rand(length(even_numbers)) + 1,
    element_to_check_1: even_numbers[e_check_index]
) else (
    element_to_check_1: "none"
);

var_even: even_numbers;
var_member_even: if member(element_to_check_1, even_numbers) then
    "\\(" + string(element_to_check_1) + " ∈ A\\)"
else
    "\\(" + string(element_to_check_1) + " ∉ A\\)";

/* Multiple choice options */
ta_1: random_permutation([["x ∉ A", false], ["x ∈ A", true], ["x ㄷ A", false]]);
ta_2: random_permutation([["x ∉ A", true], ["x ∈ A", false], ["x ㄷ A", false]]);

ta_3_a: sconcat(element_to_check, " ∈ A ");
ta_3_b: sconcat(element_to_check_1, " ∉ A ");
ta_3_c: sconcat("B ㄷ A");
ta_3_d: sconcat("A ㄷ B");
ta_3: random_permutation([[ta_3_a, true], [ta_3_b, true], [ta_3_c, false], [ta_3_d, false]]);
var_even_set:substring(string(var_even), 2, slength(string(var_even)) );</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;\({@ta_1@}\)&lt;/p&gt;
&lt;p&gt;\({@ta_2@}\)&lt;/p&gt;
&lt;p&gt;\({@ta_3@}\)&lt;/p&gt;
&lt;p&gt;\({@selected_number@}\)&lt;/p&gt;
&lt;p&gt; \({@var_even@}\)&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans2</name>
      <type>radio</type>
      <tans>ta_2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans3</name>
      <type>checkbox</type>
      <tans>ta_3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta_1)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Remember, to express that an object \( x \) is an element of set \( A \), you should use the notation \( x \in A \).&lt;/p&gt;</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;Remember, to express that an object \( x \) is not an element of set \( A \), you should use the notation \( x \notin A \).  &lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta_3)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>595633336</deployedseed>
    <deployedseed>1035473293</deployedseed>
    <deployedseed>537065522</deployedseed>
    <deployedseed>1412378921</deployedseed>
    <deployedseed>598650226</deployedseed>
    <deployedseed>365033703</deployedseed>
    <deployedseed>97484396</deployedseed>
    <deployedseed>552798541</deployedseed>
    <deployedseed>321431416</deployedseed>
    <deployedseed>1107362640</deployedseed>
    <deployedseed>1959451871</deployedseed>
    <deployedseed>1041390328</deployedseed>
    <deployedseed>220700190</deployedseed>
    <deployedseed>1566315287</deployedseed>
    <deployedseed>493966058</deployedseed>
    <deployedseed>2003844552</deployedseed>
    <deployedseed>1918575727</deployedseed>
    <deployedseed>966999930</deployedseed>
    <deployedseed>907343728</deployedseed>
    <deployedseed>1198585579</deployedseed>
    <deployedseed>2031347994</deployedseed>
    <deployedseed>478908022</deployedseed>
    <deployedseed>901822323</deployedseed>
    <deployedseed>1208481382</deployedseed>
    <deployedseed>426058151</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta_1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_2))</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>mcq_correct(ta_3)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
