<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Propositional-Logic-and-Set-Theory-(NS-and-SS)/Compound-(or-complex)-propositions/Statement-and-its-contrapositive,-truth-values-Q4.xml">
  <question type="stack">
    <name>
      <text>Statement and its contrapositive, truth values Q4</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Give the contrapositive of the given statement, then their truth values.&lt;/p&gt;
&lt;hr /&gt;
&lt;p&gt;&lt;strong&gt;Statement:&lt;/strong&gt; "\(2+2=4\) only if \(1+1=3.\)"&lt;/p&gt;
&lt;hr /&gt;
&lt;p&gt;Contrapositive: [[input:ans1]][[validation:ans1]][[feedback:prt1]]&lt;/p&gt;
&lt;p&gt;The given statement is [[input:ans2]], and its contrapositive is [[input:ans3]][[feedback:prt2]]&lt;/p&gt;
&lt;p&gt;[[validation:ans2]] [[validation:ans3]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;&lt;strong&gt;Worked solution&lt;/strong&gt;&lt;/h3&gt;
&lt;hr /&gt;
&lt;p&gt;&lt;strong&gt;Statement:&lt;/strong&gt; "\(2+2=4\) only if \(1+1=3\)."&lt;/p&gt;
&lt;hr /&gt;
&lt;p&gt;We recall that the contrapositive of \(P \implies Q\), is \(\neg Q \implies \neg P\). &lt;/p&gt;
&lt;p&gt;Contrapositive: "1+1=2 is a sufficient condition for 2+2=5."&lt;/p&gt;
&lt;p&gt;The given statement is false, and its contrapositive is false.&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>ta1: random_permutation([["1+1=2 is a sufficient condition for  2+2=5.", true],["1+1=2 is a necessary condition for 2+2=5.", false],["2+2=5 implies 1+1=2.", false],["Whenever 1+1=3,  2+2=4.", false]]);
ta2: false;
ta3: false;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;&lt;strong&gt;Statement:&lt;/strong&gt; "\(2+2=4\) only if \(1+1=3\)" &lt;/p&gt;&lt;p&gt;&lt;strong&gt;Answer:&lt;/strong&gt; {@mcq_correct(ta1)[1]@}&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text>&lt;p id="author"&gt;Mary Sayuni&lt;/p&gt;
&lt;p id="concept"&gt;N/A&lt;/p&gt;
&lt;p id="reviewer"&gt;Nixson Kiplagat&lt;/p&gt;
&lt;p id="reviewer"&gt;Wastalas Motognon&lt;/p&gt;
&lt;p id="description"&gt;Give the inverse of the given statement, then their truth values.&lt;/p&gt;</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;p&gt;Correct answer, well done.&lt;/p&gt;</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;&lt;span style="font-size: 1.5em; color: orange;"&gt; Your answer is partially correct.&lt;/span&gt;&lt;/p&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;p&gt;Incorrect answer.&lt;/p&gt;</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>boolean</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>boolean</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta1)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1189135307</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
