<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Analytic-Geometry-(NS)/Intersection-of-a-circle-with-a-line-and-tangent-line-to-a-circl/Tangent-Line-and-Circle-with-x-axis-Center.xml">
  <question type="stack">
    <name>
      <text>Tangent Line and Circle with x-axis Center</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Suppose the line \( y = {@slope*x@} \) is tangent to a circle at point \( P({@p_x@},{@p_y@}) \). 
&lt;br&gt;If the center of the circle is on the \(x\)-axis, then find the following.&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(a) The center of the circle:&lt;/p&gt;
&lt;p&gt; Center: [[input:ans1]] [[validation:ans1]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt1]]&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(b) The radius of the circle:&lt;/p&gt;
&lt;p&gt;Radius: [[input:ans2]] [[validation:ans2]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt2]]&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;(b) The equation of the circle:&lt;/p&gt;
&lt;p&gt;Equation: [[input:ans3]] [[validation:ans3]]&lt;/p&gt;
&lt;p&gt;[[feedback:prt3]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked Solution&lt;/h3&gt;
&lt;p&gt; Let the center of the circle be \(C(h,0)\). We need to find the value of \(h\).
&lt;br&gt;Since the line \( y={@slope*x@} \) is tangent to the circle at point \( P({@p_x@},{@p_y@}) \), the tangent line is perpendicular to the radius at point \(P\). 
&lt;br&gt;This means the slope of line \(CP\) is the negative reciprocal of the slope of \( y={@slope@} x \).&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;The slope of the line \( y= {@slope*x@} \) is \({@slope@}\). Therefore, the slope of line \(CP\) is \({@slope1@} \).&lt;/p&gt;
&lt;p&gt;Using the slope formula: &lt;/p&gt;
&lt;p style="margin-left:30px"&gt;The slope of \( CP = \frac{{@p_y@} - 0}{{@p_x@} - h }= {@slope1@} \).&lt;/p&gt;
&lt;p&gt;Solve for \(h\):&lt;/p&gt;
&lt;p style="margin-left:30px"&gt;\(\frac{{@p_y@}}{{@p_x@} - h }= {@slope1@}\)&lt;/p&gt;
&lt;p style="margin-left:30px"&gt;\({@p_y@}= {@(p_x - h)*(slope1)@}\)&lt;/p&gt;
&lt;p style="margin-left:30px"&gt;\(h ={@-((p_y/(slope1) )-p_x)@}\)&lt;/p&gt; 
&lt;p&gt;So, the center of the circle is \(C({@h_val@},0)\).&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;The distance between the center \(C\) and point \(P\) is the radius \(r\) of the circle. &lt;br&gt;Using the distance formula:&lt;/p&gt;
&lt;p style="margin-left:30px"&gt;\(\begin{align}
\displaystyle r &amp; = \sqrt{(x_1-h)^2+(y_1-k)^2} \\
&amp;= \sqrt{({@p_x@} {@if c_x&lt;0 then "+" else "-" @} {@abs(c_x)@})^2 + ({@p_y@} - {@c_y@})^2}  \\
&amp;=  \sqrt{({@p_x-c_x@})^2 + ({@p_y -c_y@})^2}  \\
&amp;= \sqrt{{@(p_x-c_x)^2 + (p_y -c_y)^2@}} \\
&amp;={@sqrt((p_x-c_x)^2 + (p_y -c_y)^2)@}
\end{align}\)&lt;/p&gt;
&lt;p&gt;Now we have the center \(C({@c_x@},{@c_y@})\) and the radius \(r = {@ta2@}\).&lt;/p&gt;

&lt;br&gt;

&lt;p&gt;The equation of a circle centered at \((h,k)\) with radius \(r\) is given by:&lt;/p&gt;
&lt;p style="margin-left:30px"&gt;\( (x - h)^2 + (y - k)^2 = r^2 \)&lt;/p&gt;
&lt;p&gt;Plugging in the values we found, gives the equation of the circle:&lt;/p&gt;
&lt;p style="margin-left:30px"&gt;\((x {@if c_x&lt;0 then "+" else "-" @} {@abs(c_x)@})^2 + y^2 = {@ta2^2@}\)&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Question Variables*/
slope:rand_with_prohib(-2,2,[0]);
p_x:rand_with_prohib(-2,2,[0]);
p_y:slope*p_x;
slope1:-1/slope;
c_x:-((p_y/slope1)-p_x);
c_y:0;
h_val:c_x;
ta1:[h_val,0];
ta2:ratsimp(sqrt((p_x-h_val)^2+(p_y-0)^2));
ta3:(x-h_val)^2+(y-0)^2=ta2^2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;The line \( y = {@slope*x@} \) is tangent to a circle at point \( P({@p_x@},{@p_y@}) \)&lt;/p&gt;
&lt;p&gt;Center: \({@ta1@}\)&lt;/p&gt;
&lt;p&gt;Radius: \({@ta2@}\)&lt;/p&gt;
&lt;p&gt;Equation: \({@ta3@}\)&lt;/p&gt;</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[?,?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1: Check for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1: Check for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1: Check for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1097653564</deployedseed>
    <deployedseed>1961864816</deployedseed>
    <deployedseed>1859365702</deployedseed>
    <deployedseed>2058867260</deployedseed>
    <deployedseed>1254668605</deployedseed>
    <deployedseed>67360108</deployedseed>
    <deployedseed>1357067007</deployedseed>
    <deployedseed>1336609810</deployedseed>
    <deployedseed>884489855</deployedseed>
    <deployedseed>183943631</deployedseed>
    <deployedseed>1436773572</deployedseed>
    <deployedseed>1393330079</deployedseed>
    <deployedseed>1068237912</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
            
