<?xml version="1.0" encoding="UTF-8"?>

                <quiz xmlns="http://stack-assessment.org/2025/moodle-question" xmlns:xi="http://www.w3.org/2001/XInclude" xml:base="../../../../stack/top/Mathematics-for-NS-SS-23-24-Question-Bank/Propositional-Logic-and-Set-Theory-(NS-and-SS)/Logical-connectives/Truth-Tables-and-Logical-Connectives-Interchanging-Operators-an.xml">
<!-- question: 222434  -->
  <question type="stack">
    <name>
      <text>Truth Tables and Logical Connectives: Interchanging Operators and Analyzing Patterns (narrow formatting)</text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;Logical connectives can be used together to create more complex statements.&lt;/p&gt;
&lt;hr&gt;
&lt;p&gt;Let \(p\) and \(q\) be statements.&lt;/p&gt;
&lt;p&gt;1. (a) Complete the truth tables of \( (\neg p \vee q) \wedge q\) and that of \( (\neg p \wedge q) \vee q\)&lt;/p&gt;
&lt;table border="1"&gt;
    &lt;tbody&gt;
        &lt;tr&gt;
            &lt;th&gt;p &lt;/th&gt;
            &lt;th&gt; q&lt;/th&gt;
            &lt;th&gt;\( \neg p \)&lt;/th&gt;
            &lt;th&gt;\(\neg p \vee q\)&lt;br&gt;&lt;br&gt;&lt;/th&gt;
            &lt;th&gt; \( (\neg p \vee q) \wedge q\) &lt;/th&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt; True &lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans1ss]]&lt;/td&gt;
            &lt;td&gt;[[input:ans1s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans1]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;False &lt;/td&gt;
            &lt;td&gt;[[input:ans3ss]]&lt;/td&gt;
            &lt;td&gt;[[input:ans3s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans3]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans5ss]]&lt;/td&gt;
            &lt;td&gt;[[input:ans5s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans5]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt; False &lt;/td&gt;
            &lt;td&gt;[[input:ans7ss]]&lt;/td&gt;
            &lt;td&gt;[[input:ans7s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans7]]&lt;/td&gt;
        &lt;/tr&gt;
    &lt;/tbody&gt;
&lt;/table&gt;

&lt;table border="1"&gt;
    &lt;tbody&gt;
        &lt;tr&gt;
            &lt;th&gt;p &lt;/th&gt;
            &lt;th&gt; q&lt;/th&gt;
            &lt;th&gt;\( \neg p \wedge q\)&lt;/th&gt;
            &lt;th&gt; \( (\neg p \wedge q) \vee q\)&lt;/th&gt;
            &lt;th&gt; &lt;/th&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt; True &lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans2s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans2]] &lt;/td&gt;
            &lt;td&gt;[[feedback:prt1]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;False &lt;/td&gt;
            &lt;td&gt;[[input:ans4s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans4]] &lt;/td&gt;
            &lt;td&gt;[[feedback:prt2]]&lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans6s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans6]] &lt;/td&gt;
            &lt;td&gt;[[feedback:prt3]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt; False &lt;/td&gt;
            &lt;td&gt;[[input:ans8s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans8]]&lt;/td&gt;
            &lt;td&gt;[[feedback:prt4]]&lt;/td&gt;
        &lt;/tr&gt;
    &lt;/tbody&gt;
&lt;/table&gt;
&lt;p&gt;(b) What do you notice about the truth tables of \( (\neg p \vee q) \wedge q\) and that of \( (\neg p \wedge q) \vee q\)? [[input:ans9]] [[feedback:prt5]] [[validation:ans9]]&lt;/p&gt;
&lt;p&gt;Based on this result, a creative student concludes that you can always interchange \( \vee \) and \( \wedge \) without changing the truth table.&lt;/p&gt;
&lt;p&gt;Is the student right? [[input:ans10]] [[validation:ans10]] [[feedback:prt6]]&lt;/p&gt;
&lt;p&gt;[[validation:ans1s]] [[validation:ans1]] [[validation:ans2]][[validation:ans3s]] [[validation:ans3]] [[validation:ans4]][[validation:ans5s]] [[validation:ans5]]
[[validation:ans6]] [[validation:ans7s]][[validation:ans7]][[validation:ans8]] [[validation:ans1ss]] [[validation:ans2s]] [[validation:ans3ss]] [[validation:ans4s]] [[validation:ans5ss]] [[validation:ans6s]] [[validation:ans7ss]] [[validation:ans8s]]&lt;/p&gt;
&lt;hr&gt;
&lt;p&gt;2. (a) Complete the truth tables of \( (\neg p \vee q) \wedge p\) and that of \( (\neg p \wedge q) \vee p\)&lt;/p&gt;
&lt;table border="1"&gt;
    &lt;tbody&gt;
        &lt;tr&gt;
            &lt;th&gt; p &lt;/th&gt;
            &lt;th&gt; q &lt;/th&gt;
            &lt;th&gt;\( \neg p \)&lt;/th&gt;
            &lt;th&gt;\(\neg p \vee q\)&lt;/th&gt;
            &lt;th&gt; \( (\neg p \vee q) \wedge p\) &lt;/th&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt; True&lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans11ss]]&lt;/td&gt;
            &lt;td&gt;[[input:ans11s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans11]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt; False &lt;/td&gt;
            &lt;td&gt;[[input:ans13ss]]&lt;/td&gt;
            &lt;td&gt;[[input:ans13s]]&lt;/td&gt;
            &lt;td&gt; [[input:ans13]]&lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans15ss]]&lt;/td&gt;
            &lt;td&gt;[[input:ans15s]]&lt;/td&gt;
            &lt;td&gt; [[input:ans15]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans17ss]]&lt;/td&gt;
            &lt;td&gt;[[input:ans17s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans17]] &lt;/td&gt;
        &lt;/tr&gt;
    &lt;/tbody&gt;
&lt;/table&gt;

&lt;table border="1"&gt;
    &lt;tbody&gt;
        &lt;tr&gt;
            &lt;th&gt; p &lt;/th&gt;
            &lt;th&gt; q &lt;/th&gt;
            &lt;th&gt;&amp;nbsp;\( \neg p \wedge q \)&lt;/th&gt;
            &lt;th&gt; \( (\neg p \wedge q) \vee p\)&lt;/th&gt;
            &lt;th&gt;&lt;/th&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt; True&lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans12s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans12]] &lt;/td&gt;
            &lt;td&gt;[[feedback:prt7]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt; False &lt;/td&gt;
            &lt;td&gt;[[input:ans14s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans14]] &lt;/td&gt;
            &lt;td&gt;[[feedback:prt8]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans16s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans16]] &lt;/td&gt;
            &lt;td&gt;[[feedback:prt9]] &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;[[input:ans18s]]&lt;/td&gt;
            &lt;td&gt;[[input:ans18]] &lt;/td&gt;
            &lt;td&gt;[[feedback:prt10]] &lt;/td&gt;
        &lt;/tr&gt;
    &lt;/tbody&gt;
&lt;/table&gt;
&lt;p&gt;(b) What do you notice about the truth tables of \( (\neg p \vee q) \wedge p\) and that of \( (\neg p \wedge q) \vee p\)? [[input:ans19]] [[feedback:prt11]] [[validation:ans19]]&lt;/p&gt;
&lt;p&gt;(c) What do you think of the rule formulated in 1.(a) by the student? [[input:ans20]] [[feedback:prt12]]&lt;/p&gt;
&lt;p&gt;[[validation:ans20]] [[validation:ans11]] [[validation:ans12]] [[validation:ans13]] [[validation:ans14]][[validation:ans15]] [[validation:ans16]]
[[validation:ans17]] [[validation:ans18]]&amp;nbsp;[[validation:ans11ss]] [[validation:ans11s]] [[validation:ans12s]][[validation:ans13ss]] [[validation:ans13s]] [[validation:ans14s]] [[validation:ans15ss]] [[validation:ans15s]][[validation:ans16s]] [[validation:ans17ss]]
[[validation:ans17s]] [[validation:ans18s]]&lt;/p&gt;</text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt; Worked solution&lt;/h3&gt;

&lt;p&gt;1. a Complete the truth tables of \( (\neg p \vee q) \wedge q\) and that of \( (\neg p \wedge q) \vee q\)&lt;/p&gt;
&lt;table border="1"&gt;
    &lt;tbody&gt;
        &lt;tr&gt;
            &lt;th&gt;p &lt;/th&gt;
            &lt;th&gt; q&lt;/th&gt;
            &lt;th&gt;\( \neg p \)&lt;/th&gt;
            &lt;th&gt;\(\neg p \vee q\)&lt;br&gt;&lt;br&gt;&lt;/th&gt;
            &lt;th&gt; \( (\neg p \vee q) \wedge q\) &lt;/th&gt;
            &lt;th&gt;\( \neg p \wedge q\)&lt;/th&gt;
            &lt;th&gt; \( (\neg p \wedge q) \vee q\)&lt;/th&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt; True &lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;False &lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@} &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt; False &lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
        &lt;/tr&gt;
    &lt;/tbody&gt;
&lt;/table&gt;
&lt;p&gt; 1.(b). We notice that the tables are the same. So the student is right by making such claiming.&lt;/p&gt;
&lt;hr&gt;
&lt;p&gt;2. (a) Complete the truth tables of \( (\neg p \vee q) \wedge p\) and that of \( (\neg p \wedge q) \vee p\)&lt;/p&gt;
&lt;table border="1"&gt;
    &lt;tbody&gt;
        &lt;tr&gt;
            &lt;th&gt;p &lt;/th&gt;
            &lt;th&gt; q&lt;/th&gt;
            &lt;th&gt;\( \neg p \)&lt;/th&gt;
            &lt;th&gt;\(\neg p \vee q\)&lt;br&gt;&lt;br&gt;&lt;/th&gt;
            &lt;th&gt; \( (\neg p \vee q) \wedge p\) &lt;/th&gt;
            &lt;th&gt;\( \neg p \wedge q\)&lt;/th&gt;
            &lt;th&gt; \( (\neg p \wedge q) \vee p\)&lt;/th&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt; True&lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt; False &lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@tat@} &lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;True&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
        &lt;/tr&gt;
        &lt;tr&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;False&lt;br&gt;&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@tat@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
            &lt;td&gt;{@taf@}&lt;/td&gt;
        &lt;/tr&gt;
    &lt;/tbody&gt;
&lt;/table&gt;
&lt;/p&gt;
&lt;p&gt; 2.(b). We notice in the second case that the truth tables are not the same. Therefore, the student's guess is false.&lt;/p&gt;</text>
    </generalfeedback>
    <defaultgrade>5</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>tat: true
taf: false
Q1: not p;
Q2: (not p) nounor q;
Q3: (not p) nounand q;
P1: (( not p) nounor q) nounand q;
P2: (( not p) nounand q) nounor q;
P3: (( not p) nounor q) nounand p;
P4: (( not p) nounand q) nounor p;
ta5: random_permutation([["The tables are the same", true],["The tables are not the same", false]]); 
ta6: false;
ta6: true;
ta11: random_permutation([["The tables are the same", false],["The tables are not the same", true]]); 
ta12: random_permutation([["The student's guess is true", false],["The student's guess is false", true]]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>How to fill in the truth tables for composed statements?</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt; Your answer is partially correct.&lt;/span&gt;</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>symbol</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>boolean</type>
      <tans>ta6</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11s</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11ss</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12s</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13s</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13ss</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans14</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans14s</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans15</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans15s</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans15ss</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans16</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans16s</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans17</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans17s</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans17ss</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans18</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans18s</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans19</name>
      <type>dropdown</type>
      <tans>ta11</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1s</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1ss</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans20</name>
      <type>dropdown</type>
      <tans>ta12</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2s</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3s</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3ss</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4s</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5s</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5ss</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6s</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7s</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7ss</name>
      <type>boolean</type>
      <tans>tat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8s</name>
      <type>boolean</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>dropdown</type>
      <tans>ta5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:*/
cor_ans1:is(ans1ss=false and ans1s=true and ans1=true);
/*Node 2:*/
cor_ans2: is(ans2s=false and ans2=true);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans1</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/2</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans2</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/2</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:*/
cor_ans17:is(ans17ss=true and ans17s=true and ans17=false);
/*Node 2:*/
cor_ans18: is(ans18s=false and ans18=false);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans17</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans18</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>0.2500000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans19</sans>
        <tans>mcq_correct(ta11)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt12</name>
      <value>0.2500000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans20</sans>
        <tans>mcq_correct(ta12)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt12-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:*/
cor_ans3:is(ans3ss=false and ans3s=false and ans3=false);
/*Node 2:*/
cor_ans4: is(ans4s=false and ans4=false);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>cor_ans3</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>cor_ans4</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:*/
cor_ans5:is(ans5ss=true and ans5s=true and ans5=true);
/*Node 2:*/
cor_ans6: is(ans6s=true and ans6=true);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans5</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans6</sans>
        <tans>tat</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:*/
cor_ans7:is(ans7ss=true and ans7s=true and ans7=false);
/*Node 2:*/
cor_ans8: is(ans8s=false and ans8=false);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans7</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans8</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>0.2500000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>mcq_correct(ta5)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>0.2500000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>ans10</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:*/
cor_ans11:is(ans11ss=false and ans11s=true and ans11=true);
/*Node 2:*/
cor_ans12: is(ans12s=false and ans12=true);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans11</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans12</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:*/
cor_ans13:is(ans13ss=false and ans13s=false and ans13=false);
/*Node 2:*/
cor_ans14: is(ans14s=false and ans14=true);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans13</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans14</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:*/
cor_ans15:is(ans15ss=true and ans15s=true and ans15=false);
/*Node 2:*/
cor_ans16: is(ans16s=true and ans16=true);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans15</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>PropLogic</answertest>
        <sans>cor_ans16</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>892362051</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans11s</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans11ss</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans12s</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans13</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans13s</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans13ss</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans14</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans14s</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans15</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans15s</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans15ss</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans16</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans16s</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans17</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans17s</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans17ss</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans18</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans18s</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans19</name>
        <value>first(mcq_correct(ta11))</value>
      </testinput>
      <testinput>
        <name>ans1s</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans1ss</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans20</name>
        <value>first(mcq_correct(ta12))</value>
      </testinput>
      <testinput>
        <name>ans2s</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans3s</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans3ss</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans4s</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans5s</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans5ss</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans6s</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans7s</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans7ss</name>
        <value>tat</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans8s</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>first(mcq_correct(ta5))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt12-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>
            
